```python
# Inicialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Trabalho de Casa 2: Arrays e Tabelas

Por favor, complete este notebook preenchendo as células fornecidas. Antes de começar, execute a célula anterior para carregar os testes fornecidos.

```python
# Execute esta célula e ignore a saída. Ela está verificando se você executou a primeira célula!
# Se você receber um NameError, por favor, execute a célula no topo deste notebook!

grader
```

**Recurso útil:**
- [Referência Python](http://data8.org/sp22/python-reference.html): Folha de dicas de métodos úteis de array e tabela usados em Data 8!

**Leituras recomendadas:**
- [Arrays](https://inferentialthinking.com/chapters/05/1/Arrays.html)
- [O que é Ciência de Dados?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causalidade e Experimentos](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programando em Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

Para todos os problemas que você deve escrever explicações e frases, você **deve** fornecer sua resposta no espaço designado. Além disso, ao longo deste trabalho de casa e todos os futuros, **por favor, certifique-se de não reatribuir variáveis ao longo do notebook!** Por exemplo, se você usar `max_temperature` em sua resposta a uma pergunta, não a reatribua mais tarde. Caso contrário, você falhará em testes que pensou que estava passando anteriormente!

**Nota: Este trabalho de casa tem testes ocultos. Isso significa que mesmo que os testes digam que 100% passaram, isso não significa que sua nota final será 100%. Estaremos executando mais testes de correção depois que todos entregarem o trabalho de casa.**

Compartilhar respostas diretamente não é permitido, mas discutir problemas com a equipe do curso ou com outros alunos é incentivado. Consulte a página de políticas para saber mais sobre como aprender de forma cooperativa.

Você deve começar cedo para ter tempo de obter ajuda se estiver preso.

## 1. Criando Arrays

```python
# Execute esta célula para configurar o notebook, mas por favor, não a altere.

import numpy as np
from datascience import *
import d8error
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

**Questão 1.** Crie um array chamado `weird_numbers` contendo os seguintes números (na ordem dada) **(4 Pontos)**:

1. -2
2. o piso de 12.6
3. 3
4. 5 elevado ao teto de 5.3

*Dica:* `floor` e `ceil` são funções do módulo `math`. A importação de módulos é abordada no Laboratório 2!

*Nota:* As listas Python são diferentes/se comportam de maneira diferente dos arrays NumPy. Em Data 8, usamos arrays NumPy, então por favor crie um **array**, não uma lista Python.

```python
# Nossa solução envolveu uma linha extra de código antes de criar
# weird_numbers.
...
weird_numbers = ...
weird_numbers
```

```python
grader.check("q1_1")
```

**Questão 2.** Crie um array chamado `book_title_words` contendo as seguintes três strings: "Eats", "Shoots" e "Leaves". **(4 Pontos)**

```python
book_title_words = ...
book_title_words
```

```python
grader.check("q1_2")
```

As strings têm um método chamado `join`. `join` recebe um argumento, um array de strings. Ele retorna uma única string. Especificamente, o valor de `a_string.join(an_array)` é uma única string que é a [concatenação](https://en.wikipedia.org/wiki/Concatenation) ("juntando") de todas as strings em `an_array`, **exceto** `a_string` é inserida entre cada string.

**Questão 3.** Use o array `book_title_words` e o método `join` para fazer duas strings **(4 Pontos)**:

1. "Eats, Shoots, and Leaves" (chame esta de `with_commas`)
2. "Eats Shoots and Leaves" (chame esta de `without_commas`)

*Dica:* Se você não tem certeza do que `join` faz, primeiro tente apenas chamar, por exemplo, `"data8".join(book_title_words)`.

```python
with_commas = ...
without_commas = ...

# Estas linhas são fornecidas apenas para imprimir suas respostas.
print('with_commas:', with_commas) 
print('without_commas:', without_commas)
```

```python
grader.check("q1_3")
```

## 2. Indexando Arrays

Estes exercícios oferecem prática no acesso a elementos individuais de arrays. Em Python (e em muitas linguagens de programação), cada elemento é acessado por seu *índice*; por exemplo, o primeiro elemento é o elemento no índice 0. Os índices devem ser **inteiros**.

***Nota:* Se você tem experiência prévia em codificação, pode estar familiarizado com a notação de colchetes. NÃO use a notação de colchetes ao indexar (ou seja, `arr[0]`), pois isso pode produzir diferentes tipos de dados de saída do que o que estaremos esperando. Isso pode fazer com que você falhe em um teste do corretor automático.**

Certifique-se de consultar a [Referência Python](http://data8.org/sp22/python-reference.html) no site se você se sentir preso!

**Questão 1.** A célula abaixo cria um array de alguns números. Defina `third_element` como o terceiro elemento de `some_numbers`. **(4 Pontos)**

```python
some_numbers = make_array(-1, -3, -6, -10, -15)

third_element = ...
third_element
```

```python
grader.check("q2_1")
```

**Questão 2.** A próxima célula cria uma tabela que exibe algumas informações sobre os elementos de `some_numbers` e sua ordem. Execute a célula para ver a tabela parcialmente preenchida, em seguida, preencha as informações em falta (as células que dizem "Ellipsis") atribuindo `blank_a`, `blank_b`, `blank_c`, e `blank_d` aos elementos corretos na tabela. **(4 Pontos)**

*Dica:* Substitua os `...` por strings ou números. Como lembrete, os índices devem ser **inteiros**.

```python
blank_a = ...
blank_b = ...
blank_c = ...
blank_d = ...
elements_of_some_numbers = Table().with_columns(
    "English name for position", make_array("first", "second", blank_a, blank_b, "fifth"),
    "Index",                     make_array(blank_c, 1, 2, blank_d, 4),
    "Element",                   some_numbers)
elements_of_some_numbers
```

```python
grader.check("q2_2")
```

**Questão 3.** Às vezes, você vai querer encontrar o **último** elemento de um array. Suponha que um array tenha 142 elementos. Qual é o índice do último elemento? **(4 Pontos)**

```python
index_of_last_element = ...
```

```python
grader.check("q2_3")
```

Muitas vezes, você não sabe o número de elementos em um array, seu *comprimento*. (Por exemplo, pode ser um grande conjunto de dados que você encontrou na Internet.) A função `len` recebe um único argumento, um array, e retorna um número inteiro que representa o `comprimento` desse array.

**Questão 4.** A célula abaixo carrega um array chamado `president_birth_years`. Chamar `tbl.column(...)` em uma tabela retorna um array da coluna especificada, neste caso a coluna `Birth Year` da tabela `president_births`. O último elemento desse array é o mais recente entre os anos de nascimento de todos os presidentes falecidos. Atribua esse ano a `most_recent_birth_year`. **(4 Pontos)**

```python
president_birth_years = Table.read_table("president_births.csv").column('Birth Year')

most_recent_birth_year = ...
most_recent_birth_year
```

```python
grader.check("q2_4")
```

**Questão 5.** Finalmente, atribua `min_of_birth_years` ao mínimo entre os anos de nascimento do primeiro, décimo sexto e último presidente listados em `president_birth_years`. **(4 Pontos)**

```python
min_of_birth_years = ...
min_of_birth_years
```

```python
grader.check("q2_5")
```

## 3. Aritmética Básica de Arrays

**Questão 1.** Multiplique os números 42, -4224, 424224242 e 250 por 157. Atribua cada variável abaixo de forma que `first_product` seja atribuído ao resultado de $42 * 157$, `second_product` seja atribuído ao resultado de $-4224 * 157$, e assim por diante. **(4 Pontos)**

*Nota*: Para esta questão, **não** use arrays.

```python
first_product = ...
second_product = ...
third_product = ...
fourth_product = ...
print("First Product:", first_product)
print("Second Product:", second_product)
print("Third Product:", third_product)
print("Fourth Product:", fourth_product)
```

```python
grader.check("q3_1")
```

**Questão 2.** Agora, faça o mesmo cálculo, mas usando um array chamado `numbers` e apenas um único operador de multiplicação (`*`). Armazene os 4 resultados em um array chamado `products`. **(4 Pontos)**

```python
numbers = ...
products = ...
products
```

```python
grader.check("q3_2")
```

**Questão 3.** Ops, cometemos um erro de digitação! Em vez de 157, queríamos multiplicar cada número por 1577. Calcule os produtos corretos na célula abaixo usando aritmética de array. Observe que seu trabalho é realmente fácil se você definiu anteriormente um array contendo os 4 números. **(4 Pontos)**

```python
correct_products = ...
correct_products
```

```python
grader.check("q3_3")
```

**Questão 4.** Carregamos um array de temperaturas na próxima célula. Cada número é a temperatura mais alta observada em um dia em uma estação de observação climática, principalmente dos EUA. Como são da agência governamental dos EUA [NOAA](https://www.noaa.gov/), todas as temperaturas estão em Fahrenheit.

Converta todas as temperaturas para Celsius subtraindo primeiro 32 delas, depois multiplicando os resultados por $\frac{5}{9}$, ou seja, $C = (F - 32) * \frac{5}{9}$. Após converter as temperaturas para Celsius, certifique-se de **ARREDONDAR** o resultado final para o inteiro mais próximo usando a função `np.round`. **(4 Pontos)**

```python
max_temperatures = Table.read_table("temperatures.csv").column("Daily Max Temperature")

celsius_max_temperatures = ...
celsius_max_temperatures
```

```python
grader.check("q3_4")
```

**Questão 5.** A célula abaixo carrega todas as *menores* temperaturas de cada dia (em Fahrenheit). Calcule a variação de temperatura diária para cada dia. Ou seja, calcule a diferença entre cada temperatura máxima diária e a correspondente temperatura mínima diária. **Preste atenção nas unidades e dê sua resposta em Celsius!** Certifique-se de **NÃO** arredondar sua resposta para esta questão! **(4 Pontos)**

*Nota:* Lembre-se de que na Questão 4, `celsius_max_temperatures` foi arredondado, então você pode não querer usar essa variável nesta questão.

```python
min_temperatures = Table.read_table("temperatures.csv").column("Daily Min Temperature")

celsius_temperature_ranges = ...
celsius_temperature_ranges
```

```python
grader.check("q3_5")
```

## 4. Velho Fiel

Old Faithful is a geyser in Yellowstone that erupts every 44 to 125 minutes (according to [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). People are [often told that the geyser erupts every hour](http://yellowstone.net/geysers/old-faithful/), but in fact the waiting time between eruptions is more variable. Let's take a look.

O Velho Fiel é um gêiser em Yellowstone que entra em erupção a cada 44 a 125 minutos (de acordo com a [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). As pessoas são [frequentemente informadas de que o gêiser entra em erupção a cada hora](http://yellowstone.net/geysers/old-faithful/), mas na verdade o tempo de espera entre as erupções é mais variável. Vamos dar uma olhada.

**Questão 1.** A primeira linha abaixo atribui `waiting_times` a um array de 272 tempos de espera consecutivos entre erupções, retirados de um clássico conjunto de dados de 1938. Atribua os nomes `shortest`, `longest` e `average` para que a instrução `print` esteja correta. **(4 Pontos)**

```python
waiting_times = Table.read_table('old_faithful.csv').column('waiting')

shortest = ...
longest = ...
average = ...

print("Old Faithful erupts every", shortest, "to", longest, "minutes and every", average, "minutes on average.")
```

```python
grader.check("q4_1")
```

**Questão 2.** Atribua `biggest_decrease` à maior diminuição no tempo de espera entre duas erupções consecutivas. Por exemplo, a terceira erupção ocorreu após 74 minutos e a quarta após 62 minutos, então a diminuição no tempo de espera foi de 74 - 62 = 12 minutos. **(4 Pontos)**

*Dica*: Queremos retornar o valor absoluto da maior diminuição.

*Nota*: `np.diff()` calcula a diferença entre valores subsequentes em um array. Por exemplo, chamar `np.diff()` no array `make_array(1, 8, 3, 5)` avalia para `array([8 - 1, 3 - 8, 5 - 3])`, ou `array([7, -5, 2])`.

```python
# np.diff() calcula a diferença entre valores subsequentes
# em um array NumPy.
differences = np.diff(waiting_times)
biggest_decrease = ...
biggest_decrease
```

```python
grader.check("q4_2")
```

**Questão 3.** A tabela `faithful_with_eruption_nums` contém duas colunas: `eruption_number`, que representa o número daquela erupção, e `waiting`, que representa o tempo gasto esperando após aquela erupção. Por exemplo, pegue as duas primeiras linhas da tabela:

| número da erupção | espera |
|-------------------|--------|
| 1                 | 79     |
| 2                 | 54     |

Podemos ler isso da seguinte forma: após a primeira erupção, esperamos 79 minutos pela segunda erupção. Então, após a segunda erupção, esperamos 54 minutos pela terceira erupção.

Suponha que Oscar e Wendy começaram a assistir ao Velho Fiel no início da primeira erupção. Suponha que eles assistam até o final da décima erupção. Durante parte desse tempo, eles estarão assistindo às erupções, e pelo resto do tempo, eles estarão esperando o Velho Fiel entrar em erupção. Quantos minutos eles passarão esperando pelas erupções? **(4 Pontos)**

*Dica #1:* Uma maneira de abordar este problema é usar o método `take` ou `where` na tabela `faithful_with_eruption_nums`.

*Dica #2:* `first_nine_waiting_times` deve ser um array.

```python
# As duas linhas seguintes carregam nossa tabela faithful_with_eruption_nums
faithful = Table.read_table('old_faithful.csv').drop("eruptions")
faithful_with_eruption_nums = faithful.with_column("eruption number", np.arange(faithful.num_rows) + 1).select(1, 0)

first_nine_waiting_times = ...
total_waiting_time_until_tenth = ...
total_waiting_time_until_tenth
```

```python
grader.check("q4_3")
```

**Questão 4.** Vamos imaginar que seu palpite para o próximo tempo de espera fosse sempre apenas a duração do tempo de espera anterior. Se você sempre adivinhasse o tempo de espera anterior, qual seria o seu erro médio ao adivinhar os tempos de espera? **(4 Pontos)**

Por exemplo, como os primeiros quatro tempos de espera são 79, 54, 74 e 62, a diferença média entre seu palpite e o tempo real para apenas as segunda, terceira e quarta erupções seria $\frac{|79-54|+ |54-74|+ |74-62|}{3} = 19$.

```python
differences = np.diff(waiting_times)
average_error = ...
average_error
```

```python
grader.check("q4_4")
```

## 5. Tabelas

**Questão 1.** Suponha que você tenha 4 maçãs, 3 laranjas e 3 abacaxis. (Talvez você esteja usando Python para resolver um problema de álgebra do ensino médio.) Crie uma tabela que contenha essas informações. Ela deve ter duas colunas: `fruit name` e `count`. Atribua a nova tabela à variável `fruits`. **(4 Pontos)**

*Nota:* Use palavras em minúsculas e no singular para o nome de cada fruta, como `"apple"`.


```python
# Nossa solução usa 1 instrução dividida em 3 linhas.
frutas = ...
        ...
        ...
frutas
```

```python
grader.check("q5_1")
```

**Questão 2.** O arquivo `inventory.csv` contém informações sobre o inventário de uma banca de frutas. Cada linha representa o conteúdo de uma caixa de frutas. Carregue-o como uma tabela chamada `inventory` usando a função `Table.read_table()`. `Table.read_table(...)` recebe um argumento (nome do arquivo de dados em formato de string) e retorna uma tabela. **(4 Pontos)**

```python
inventory = ...
inventory
```

```python
grader.check("q5_2")
```

**Questão 3.** Cada caixa na banca de frutas contém uma fruta diferente? Defina `all_different` como `True` se cada caixa contiver uma fruta diferente ou como `False` se várias caixas contiverem a mesma fruta. **(4 Pontos)**

*Dica:* Você não precisa escrever código para calcular o valor Verdadeiro/Falso para `all_different`. Basta olhar para a tabela `inventory` e atribuir `all_different` a `True` ou `False` de acordo com o que você pode ver da tabela ao responder à pergunta.

```python
all_different = ...
all_different
```

```python
grader.check("q5_3")
```

**Questão 4.** O arquivo `sales.csv` contém o número de frutas vendidas de cada caixa no último sábado. Ele tem uma coluna extra chamada `price per fruit ($)` que é o preço *por item de fruta* para a fruta naquela caixa. As linhas estão na mesma ordem que a tabela `inventory`. Carregue esses dados em uma tabela chamada `sales`. **(5 Pontos)**

```python
sales = ...
sales
```

```python
grader.check("q5_4")
```

**Questão 5.** Quantas frutas a loja vendeu no total naquele dia? **(5 Pontos)**


```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("q5_5")
```

**Questão 6.** Qual foi a receita total da loja (o preço total de todas as frutas vendidas) naquele dia? **(5 Pontos)**

*Dica:* Se você estiver preso, pense primeiro em como calcularia a receita total apenas das vendas de uvas.

```python
total_revenue = ...
total_revenue
```

```python
grader.check("q5_6")
```

**Questão 7.** Crie uma nova tabela chamada `remaining_inventory`. Ela deve ter as mesmas linhas e colunas que `inventory`, exceto que a quantidade de frutas vendidas de cada caixa deve ser subtraída da contagem **original** daquela caixa, de modo que a coluna `count` seja **atualizada para ser** a quantidade de frutas restantes após o sábado. **(5 Pontos)**

```python
remaining_inventory = ...
    ...
    ...
    ...
remaining_inventory
```

```python
grader.check("q5_7")
```

Você terminou o Trabalho de Casa 2!

**Passos importantes para a submissão:**
1. Execute os testes e verifique se todos passaram.
2. Escolha **Salvar Notebook** no menu **Arquivo**, então **execute a última célula**.
3. Clique no link para baixar o arquivo zip.
4. Em seguida, envie o arquivo zip para a tarefa correspondente de acordo com as instruções do seu professor.

**É sua responsabilidade garantir que seu trabalho esteja salvo antes de executar a última célula.**

## Submissão

Certifique-se de ter executado todas as células em seu notebook em ordem antes de executar a célula abaixo, para que todas as imagens/gráficos apareçam na saída. A célula abaixo irá gerar um arquivo zip para você enviar. **Por favor, salve antes de exportar!**

```python
# Salve seu notebook primeiro, depois execute esta célula para exportar sua submissão.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 0])
```
>>> import numpy as np
>>> # It looks like you didn't make an array.
>>> type(weird_numbers) == np.ndarray
True
```
```
>>> len(weird_numbers) == 4
True
```


Tests for q1_2 (points: [0, 0, 0, 0, 4])
```
>>> import numpy as np
>>> # It looks like you didn't make an array.
>>> type(book_title_words) == np.ndarray
True
```
```
>>> # It looks like you included commas in the text.
>>> # The three pieces of text in the array should be:
>>> #   "Eats"
>>> #   "Shoots"
>>> #   "and Leaves"
>>> not any([',' in text for text in book_title_words])
True
```
```
>>> # It looks like you didn't include both words in the
>>> # last piece of text.  It should be the actual string:
>>> #   "and Leaves"
>>> 'and ' in book_title_words.item(2)
True
```
```
>>> len(book_title_words)
3
```
```
>>> book_title_words.item(0) == 'Eats' and book_title_words.item(1) == 'Shoots' and book_title_words.item(2) == 'and Leaves'
True
```


Tests for q1_3 (points: [2, 2])
```
>>> with_commas == 'Eats, Shoots, and Leaves'
True
```
```
>>> without_commas == 'Eats Shoots and Leaves'
True
```


Tests for q2_1 (points: [0, 4])
```
>>> # It looks like you wrote:
>>> #   some_numbers.item(3)
>>> # But the third element has index 2, not 3.
>>> third_element != -10
True
```
```
>>> third_element == -6
True
```


Tests for q2_2 (points: [1, 1, 1, 1])
```
>>> elements_of_some_numbers.column(0).item(2) == "third"
True
```
```
>>> elements_of_some_numbers.column(0).item(3) == "fourth"
True
```
```
>>> elements_of_some_numbers.column(1).item(0) == 0
True
```
```
>>> elements_of_some_numbers.column(1).item(3) == 3
True
```


Tests for q2_3 (points: [4])
```
>>> index_of_last_element == 141
True
```


Tests for q2_4 (points: [4])
```
>>> most_recent_birth_year == 1917
True
```


Tests for q2_5 (points: [0])
```
>>> min_of_birth_years > 0
True
```


Tests for q3_1 (points: [0])
```
>>> first_product > 0, second_product < 0, third_product > 0, fourth_product > 0
(True, True, True, True)
```


Tests for q3_2 (points: [0])
```
>>> # products should be an array!
>>> len(products) == 4
True
```


Tests for q3_3 (points: [0])
```
>>> len(correct_products) == 4
True
```


Tests for q3_4 (points: [1, 1, 1, 1])
```
>>> # It looks like you multiplied and subtracted in the wrong
>>> # order.
>>> sum(celsius_max_temperatures) != 356705.0
True
```
```
>>> sum(celsius_max_temperatures)
1280677.0
```
```
>>> len(celsius_max_temperatures)
65000
```
```
>>> celsius_max_temperatures.item(2003)
20.0
```


Tests for q3_5 (points: [0, 4, 0])
```
>>> # celsius_temperature_ranges should be an array
>>> type(celsius_temperature_ranges) is np.ndarray
True
```
```
>>> np.isclose(round(sum(celsius_temperature_ranges)), 768487)
True
```
```
>>> len(celsius_temperature_ranges)
65000
```


Tests for q4_1 (points: [0, 0])
```
>>> # Hint: shortest is a number between 40 and 50.
>>> 40 <= shortest <= 50
True
```
```
>>> # Hint: the average is between the shortest and the longest
>>> shortest <= average <= longest
True
```


Tests for q4_2 (points: [0, 0])
```
>>> # Hint: If you are getting 47 as your answer, you might be computing the biggest change 
>>> # rather than the biggest decrease!
>>> biggest_decrease == 47
False
```
```
>>> # Hint: biggest decrease is above 30, but not 47.
>>> 30 <= biggest_decrease < 47
True
```


Tests for q4_3 (points: [0, 4])
```
>>> len(first_nine_waiting_times)
9
```
```
>>> total_waiting_time_until_tenth
633
```


Tests for q4_4 (points: [0])
```
>>> 15 <= average_error <= 25
True
```


Tests for q5_1 (points: [4])
```
>>> fruits.sort(0)
fruit name | count
apple      | 4
orange     | 3
pineapple  | 3
```


Tests for q5_2 (points: [4])
```
>>> inventory.sort(0).column(0).item(0)
25274
```


Tests for q5_3 (points: [0])
```
>>> all_different in {True, False}
True
```


Tests for q5_4 (points: [5])
```
>>> sales.sort(0)
box ID | fruit name | count sold | price per fruit ($)
25274  | apple      | 0          | 0.8
26187  | strawberry | 25         | 0.15
43566  | peach      | 17         | 0.8
48800  | orange     | 35         | 0.6
52357  | strawberry | 102        | 0.25
53686  | kiwi       | 3          | 0.5
57181  | strawberry | 101        | 0.2
57930  | grape      | 355        | 0.06
```


Tests for q5_5 (points: [0])
```
>>> # We're asking for the number of *pieces* of fruit, not the
>>> # number of kinds of fruit or the number of boxes from which
>>> # there were sales.
>>> total_fruits_sold > 10
True
```


Tests for q5_6 (points: [0])
```
>>> # If you're stuck, here's a hint: You want to multiply the count
>>> # sold in each box by the per-item price of fruits in that box.
>>> # You can use elementwise multiplication for that.
>>> # Then you want the sum of those products.  Use sum().
>>> 50 <= total_revenue <= 150
True
```


Tests for q5_7 (points: [0, 2, 3])
```
>>> # It looks like your table doesn't have all 3 columns that are
>>> # in the inventory table.
>>> remaining_inventory.num_columns
3
```
```
>>> #It looks like you forgot to subtract off the sales.
>>> remaining_inventory.column("count").item(0) != 45
True
```
```
>>> remaining_inventory.where(1, 'grape')
box ID | fruit name | count
57930  | grape      | 162
```
