```python
# Inicialize o Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# Laboratório 2: Operações com Tabelas

Bem-vindo ao Laboratório 2! Esta semana, vamos aprender como importar um módulo e praticar operações com tabelas. O [Python Reference](http://data8.org/sp22/python-reference.html) tem informações que serão úteis para este laboratório.

**Leitura Recomendada**:
 * [Introdução às Tabelas](https://www.inferentialthinking.com/chapters/03/4/Introduction_to_Tables)

**Envio**: Quando terminar, execute todas as células exceto a última, selecione Arquivo > Salvar Notebook e, em seguida, execute a célula final. Depois, envie o arquivo zip baixado, que inclui o seu notebook, de acordo com as instruções do seu instrutor.

Vamos começar configurando os testes e as importações executando a célula abaixo.

```python
# Não altere esta célula; apenas execute-a.

import numpy as np
from datascience import *
import d8error
```

# 1. Revisão: Os Blocos de Construção do Código Python

Os dois blocos de construção do código Python são *expressões* e *instruções*. Uma **expressão** é um trecho de código que

* é autocontido, o que significa que faria sentido escrevê-lo em uma linha própria, e
* geralmente avalia para um valor.

Aqui estão duas expressões que ambas avaliam para 3:

    3
    5 - 2
    
Um tipo importante de expressão é a **expressão de chamada**. Uma expressão de chamada começa com o nome de uma função e é seguida pelo(s) argumento(s) dessa função entre parênteses. A função retorna algum valor, com base em seus argumentos. Algumas funções matemáticas importantes estão listadas abaixo.

| Função   | Descrição                                                      |
|----------|-----------------------------------------------------------------|
| `abs`    | Retorna o valor absoluto de seu argumento                      |
| `max`    | Retorna o máximo de todos os seus argumentos                   |
| `min`    | Retorna o mínimo de todos os seus argumentos                   |
| `pow`    | Eleva seu primeiro argumento à potência do segundo argumento   |
| `round`  | Arredonda seu argumento para o inteiro mais próximo            |

Aqui estão duas expressões de chamada que ambas avaliam para 3:

    abs(2 - 5)
    max(round(2.8), min(pow(2, 10), -1 * pow(2, 10)))

A expressão `2 - 5` e as duas expressões de chamada dadas acima são exemplos de **expressões compostas**, o que significa que são realmente combinações de várias expressões menores. `2 - 5` combina as expressões `2` e `5` por subtração. Neste caso, `2` e `5` são chamados de **subexpressões** porque são expressões que fazem parte de uma expressão maior.

Uma **instrução** é uma linha inteira de código. Algumas instruções são apenas expressões. As expressões listadas acima são exemplos.

Outras instruções *fazem algo acontecer* em vez de *ter um valor*. Por exemplo, uma **instrução de atribuição** atribui um valor a um nome.

Uma boa maneira de pensar sobre isso é que estamos **avaliando o lado direito** do sinal de igual e **atribuindo-o ao lado esquerdo**. Aqui estão algumas instruções de atribuição:

    altura = 1.3
    o_numero_cinco = abs(-5)
    diferenca_de_altura_absoluta = abs(altura - 1.688)

Uma ideia importante na programação é que coisas grandes e interessantes podem ser construídas combinando muitas coisas simples e sem importância. A chave para entender um código complicado é dividi-lo em seus componentes simples.

Por exemplo, muito está acontecendo na última instrução acima, mas na verdade é apenas uma combinação de algumas coisas. Esta imagem descreve o que está acontecendo.

![statement](statement.png)


**Questão 1.1.** Na próxima célula, atribua o nome `new_year` ao número maior entre os dois números a seguir:

* o **valor absoluto** de $2^{5}-2^{11}-2^1 + 1$, e 
* $5 \times 13 \times 31 + 7$.

Tente usar apenas uma instrução (uma linha de código). Certifique-se de verificar seu trabalho executando a célula de teste posteriormente.

```python
new_year = ...
new_year
```

```python
grader.check("q11")
```

Pedimos que você use uma linha de código na pergunta acima porque ela envolve apenas operações matemáticas. No entanto, perguntas de programação mais complicadas exigirão mais etapas. Nem sempre é uma boa ideia encaixar essas etapas em uma única linha, pois pode tornar o código mais difícil de ler e de depurar.

Uma boa prática de programação envolve dividir seu código em etapas menores e usar nomes apropriados. Você terá muita prática no resto deste curso!

# 2. Importando Código

A maioria da programação envolve trabalho muito semelhante ao que já foi feito antes. Como escrever código consome tempo, é bom contar com o código publicado por outros quando você pode. Em vez de copiar e colar, o Python nos permite **importar módulos**. Um módulo é um arquivo com código Python que definiu variáveis e funções. Ao importar um módulo, somos capazes de usar seu código em nosso próprio notebook.

Python inclui muitos módulos úteis que estão a apenas um `import` de distância. Vamos olhar para o módulo `math` como um primeiro exemplo. O módulo `math` é extremamente útil para calcular expressões matemáticas em Python.

Suponha que queremos calcular com muita precisão a área de um círculo com um raio de 5 metros. Para isso, precisamos da constante $\pi$, que é aproximadamente 3,14. Convenientemente, o módulo `math` tem `pi` definido para nós. Execute a célula seguinte para importar o módulo math:

```python
import math
radius = 5
area_of_circle = radius**2 * math.pi
area_of_circle
```

No código acima, a linha `import math` importa o módulo math. Esta instrução cria um módulo e, em seguida, atribui o nome `math` a esse módulo. Agora somos capazes de acessar quaisquer variáveis ou funções definidas dentro de `math` digitando o nome do módulo seguido por um ponto, e então seguido pelo nome da variável ou função que queremos.

    <nome do módulo>.<nome>

**Questão 2.1.** O módulo `math` também fornece o nome `e` para a base do logaritmo natural, que é aproximadamente 2,71. Calcule $e^{\pi}-\pi$, dando a ele o nome `near_twenty`.

*Lembre-se: Você também pode acessar `pi` a partir do módulo `math`!*

```python
near_twenty = ...
near_twenty
```

```python
grader.check("q21")
```

## 2.1. Acessando Funções

Na questão acima, você acessou variáveis dentro do módulo `math`.

**Módulos** também definem **funções**. Por exemplo, `math` fornece o nome `floor` para a função piso. Tendo importado `math` já, podemos escrever `math.floor(7.5)` para calcular o piso de 7.5. (Note que a função piso retorna o maior inteiro menor ou igual a um determinado número.)

**Questão 2.1.1.** Calcule o piso de pi usando `floor` e `pi` do módulo `math`. Dê o resultado o nome `floor_of_pi`.


```python
floor_of_pi = ...
floor_of_pi
```

```python
grader.check("q211")
```

Para sua referência, abaixo estão alguns exemplos adicionais de funções do módulo `math`.

Observe como diferentes funções recebem diferentes números de argumentos. Muitas vezes, a [documentação](https://docs.python.org/3/library/math.html) do módulo fornecerá informações sobre quantos argumentos são necessários para cada função.

*Dica: Se você pressionar `shift+tab` ao lado da chamada da função, a documentação dessa função aparecerá.*

```python
# Calculando logaritmos (o logaritmo de 8 na base 2).
# O resultado é 3 porque 2 elevado à potência de 3 é 8.
math.log(8, 2)
```

```python
# Calculando raízes quadradas.
math.sqrt(5)
```

Existem várias maneiras de importar e acessar código de fontes externas. O método que usamos acima - `import <nome_do_módulo>` - importa todo o módulo e requer que usemos `<nome_do_módulo>.<nome>` para acessar seu código.

Também podemos importar uma constante específica ou função em vez de todo o módulo. Observe que você não precisa usar o nome do módulo antes para referenciar aquele valor particular. No entanto, você precisa ter cuidado ao reatribuir os nomes das constantes ou funções para outros valores!

```python
# Importando apenas cos e pi de math.
# Não precisamos usar `math.` na frente de cos ou pi
from math import cos, pi
print(cos(pi))

# Mas precisamos usar na frente de outras funções de math
math.log(pi)
```

Ou podemos importar todas as funções e valores de todo o módulo.

```python
# Por último, podemos importar tudo de math usando o *
# Mais uma vez, não precisamos usar 'math.' antes
from math import *
log(pi)
```

Não se preocupe muito sobre qual tipo de importação usar. Muitas vezes, é uma escolha de estilo de codificação deixada a critério de cada programador. Neste curso, você sempre importará os módulos necessários quando executar a célula de configuração (como a primeira célula de código neste laboratório).

Vamos seguir em frente e praticar algumas das operações de tabela que você aprendeu na aula!

# 3. Operações com Tabelas

A tabela `farmers_markets.csv` contém dados sobre mercados de agricultores nos Estados Unidos (dados associados com [o USDA](https://apps.ams.usda.gov/FarmersMarketsExport/ExcelExport.aspx)). Cada linha representa um desses mercados.

Execute a próxima célula para carregar a tabela `farmers_markets`. Não haverá saída -- nenhuma saída é esperada, pois a célula contém uma instrução de atribuição. Uma instrução de atribuição não produz nenhuma saída (não gera nenhum valor).

```python
# Apenas execute esta célula

farmers_markets = Table.read_table('farmers_markets.csv')
```

**Questão 3.1.** Use o método `show` para exibir as primeiras 5 linhas de `farmers_markets`.

*Nota:* Os termos "método" e "função" tecnicamente não são a mesma coisa, mas para os propósitos deste curso, usaremos os termos de forma intercambiável.

**Dica:** `tbl.show(3)` mostrará as primeiras 3 linhas da tabela chamada `tbl`. Além disso, certifique-se de não chamar `.show()` sem um argumento, pois isso fará com que seu kernel trave!

```python
...
```

Observe que alguns dos valores nesta tabela estão faltando, conforme indicado por "nan". Isso significa que o valor não está disponível (por exemplo, se não soubermos o endereço da rua do mercado) ou não é aplicável (por exemplo, se o mercado não tiver um endereço de rua). Você também notará que a tabela tem um grande número de colunas!

### `num_columns`

A propriedade da tabela `num_columns` retorna o número de colunas em uma tabela. (Uma "propriedade" é apenas um método que não precisa ser chamado adicionando parênteses.)

Exemplo de chamada: `tbl.num_columns` retornará o número de colunas em uma tabela chamada `tbl`

**Questão 3.2.** Use `num_columns` para encontrar o número de colunas em nosso conjunto de dados de mercados de agricultores.

Atribua o número de colunas a `num_farmers_markets_columns`.


```python
num_farmers_markets_columns = ...
print("The table has", num_farmers_markets_columns, "columns in it!")
```

```python
grader.check("q32")
```

### `num_rows`

Da mesma forma, a propriedade `num_rows` informa quantas linhas estão em uma tabela.

```python
# Apenas execute esta célula

num_farmers_markets_rows = farmers_markets.num_rows
print("A tabela tem", num_farmers_markets_rows, "linhas!")
```

### `select`

A maioria das colunas é sobre produtos específicos - se o mercado vende tofu, comida para animais de estimação, etc. Se não estamos interessados nessa informação, ela apenas torna a tabela difícil de ler. Isso ocorre mais do que você imagina, porque as pessoas que coletam e publicam dados podem não saber antecipadamente o que as pessoas vão querer fazer com eles.

Em tais situações, podemos usar o método `select` da tabela para escolher apenas as colunas que queremos em uma tabela específica. Ele aceita qualquer número de argumentos. Cada um deve ser o nome de uma coluna na tabela. Ele retorna uma nova tabela com apenas essas colunas. As colunas estão na ordem *em que foram listadas como argumentos*.

Por exemplo, o valor de `farmers_markets.select("MarketName", "State")` é uma tabela com apenas o nome e o estado de cada mercado de agricultores em `farmers_markets`.

**Questão 3.3.** Use `select` para criar uma tabela com apenas o nome, cidade, estado, latitude (`y`) e longitude (`x`) de cada mercado. Chame essa nova tabela de `farmers_markets_locations`.

*Dica:* Certifique-se de ser exato ao usar nomes de colunas com `select`; verifique novamente a capitalização!


```python
farmers_markets_locations = ...
farmers_markets_locations
```

```python
grader.check("q33")
```

### `drop`

`drop` serve ao mesmo propósito que `select`, mas remove as colunas que você fornece em vez das que você não fornece. Como `select`, `drop` retorna uma nova tabela.

**Questão 3.4.** Suponha que você simplesmente não queira as colunas `FMID` e `updateTime` em `farmers_markets`. Crie uma tabela que seja uma cópia de `farmers_markets`, mas que não inclua essas colunas. Chame essa tabela de `farmers_markets_without_fmid`.

```python
farmers_markets_without_fmid = ...
farmers_markets_without_fmid
```

```python
grader.check("q34")
```

Agora, suponha que queremos responder a algumas perguntas sobre os mercados de agricultores nos EUA. Por exemplo, qual(is) mercado(s) tem a maior longitude (dada pela coluna `x`)?

Para responder a isso, vamos ordenar `farmers_markets_locations` por longitude.

```python
farmers_markets_locations.sort('x')
```

Ops, isso não respondeu à nossa pergunta porque ordenamos da menor para a maior longitude. Para olhar as maiores longitudes, teremos que ordenar em ordem inversa.

```python
farmers_markets_locations.sort('x', descending=True)
```

(O pedaço `descending=True` é chamado de *argumento opcional*. Ele tem um valor padrão de `False`, então quando você explicitamente diz à função `descending=True`, a função irá ordenar em ordem decrescente.)

### `sort`

Alguns detalhes sobre a ordenação:

1. O primeiro argumento para `sort` é o nome de uma coluna para ordenar.
2. Se a coluna contém texto, `sort` irá ordenar alfabeticamente; se a coluna contém números, ela irá ordenar numericamente - ambos em ordem crescente por padrão.
3. O valor de `farmers_markets_locations.sort("x")` é uma *cópia* de `farmers_markets_locations`; a tabela `farmers_markets_locations` não é modificada. Por exemplo, se chamássemos `farmers_markets_locations.sort("x")`, então executar `farmers_markets_locations` por si só ainda retornaria a tabela não ordenada.
4. As linhas sempre permanecem juntas quando uma tabela é ordenada. Não faria sentido ordenar apenas uma coluna e deixar as outras colunas sozinhas. Por exemplo, neste caso, se ordenássemos apenas a coluna `x`, os mercados de agricultores acabariam com as longitudes erradas.

**Questão 3.5.** Crie uma versão de `farmers_markets_locations` que seja ordenada por **latitude (`y`)**, com as maiores latitudes primeiro. Chame-a de `farmers_markets_locations_by_latitude`.


```python
farmers_markets_locations_by_latitude = ...
farmers_markets_locations_by_latitude
```

```python
grader.check("q35")
```

Agora vamos dizer que queremos uma tabela de todos os mercados de agricultores na Califórnia. A ordenação não nos ajudará muito aqui porque a Califórnia está mais próxima do meio do conjunto de dados.

Em vez disso, usamos o método de tabela `where`.

```python
california_farmers_markets = farmers_markets_locations.where('State', are.equal_to('California'))
california_farmers_markets
```

Ignore a sintaxe por um momento. Em vez disso, tente ler essa linha assim:

> Atribua o nome **`california_farmers_markets`** a uma tabela cujas linhas são as linhas na tabela **`farmers_markets_locations`** **`onde`** o **`"State"`** **`é igual a` `"California"`**.

### `where`

Agora vamos nos aprofundar um pouco mais nos detalhes. `where` recebe 2 argumentos:

1. O nome de uma coluna. `where` encontra linhas onde os valores dessa coluna atendem a algum critério.
2. Um predicado que descreve o critério que a coluna precisa atender.

O predicado no exemplo acima chamou a função `are.equal_to` com o valor desejado, 'California'. Veremos outros predicados em breve.

`where` retorna uma tabela que é uma cópia da tabela original, mas **apenas com as linhas que atendem ao predicado dado**.

**Questão 3.6.** Use `california_farmers_markets` para criar uma tabela chamada `berkeley_markets` contendo mercados de agricultores em Berkeley, Califórnia.

```python
berkeley_markets = ...
berkeley_markets
```

```python
grader.check("q36")
```

Até agora, só usamos `where` com o predicado que requer encontrar os valores em uma coluna para serem *exatamente* iguais a um determinado valor. No entanto, existem muitos outros predicados. Aqui estão alguns:

|Predicado|Exemplo|Resultado|
|-|-|-|
|`are.equal_to`|`are.equal_to(50)`|Encontra linhas com valores iguais a 50|
|`are.not_equal_to`|`are.not_equal_to(50)`|Encontra linhas com valores não iguais a 50|
|`are.above`|`are.above(50)`|Encontra linhas com valores acima (e não iguais a) 50|
|`are.above_or_equal_to`|`are.above_or_equal_to(50)`|Encontra linhas com valores acima de 50 ou iguais a 50|
|`are.below`|`are.below(50)`|Encontra linhas com valores abaixo de 50|
|`are.between`|`are.between(2, 10)`|Encontra linhas com valores acima ou iguais a 2 e abaixo de 10|
|`are.between_or_equal_to`|`are.between_or_equal_to(2, 10)`|Encontra linhas com valores acima ou iguais a 2 e abaixo ou iguais a 10|

## 4. Analisando um conjunto de dados

Agora que você está familiarizado com as operações de tabela, vamos responder a uma pergunta interessante sobre um conjunto de dados!

Execute a célula abaixo para carregar a tabela `imdb`. Ela contém informações sobre os 250 filmes mais bem avaliados no IMDb.

```python
# Apenas execute esta célula

imdb = Table.read_table('imdb.csv')
imdb
```

Frequentemente, queremos realizar várias operações - ordenação, filtragem, entre outras - para transformar uma tabela que temos em algo mais útil. Você pode fazer essas operações uma de cada vez, por exemplo:

```
primeiro_passo = tabela_original.where("col1", are.equal_to(12))
segundo_passo = primeiro_passo.sort('col2', descending=True)
```

No entanto, como o valor da expressão `tabela_original.where("col1", are.equal_to(12))` é em si uma tabela, você pode simplesmente chamar um método de tabela nela:

```
tabela_original.where("col1", are.equal_to(12)).sort('col2', descending=True)
```
Você deve organizar seu trabalho da maneira que faça mais sentido para você, usando nomes informativos para quaisquer tabelas intermediárias que você criar.

**Questão 4.1.** Crie uma tabela de filmes lançados entre 2010 e 2015 (inclusive) com classificações acima de 8. A tabela deve conter apenas as colunas `Title` e `Rating`, **nessa ordem**.

Atribua a tabela ao nome `above_eight`.

*Dica:* Pense nas etapas que você precisa realizar e tente colocá-las em uma ordem que faça sentido. Sinta-se à vontade para criar tabelas intermediárias para cada etapa, mas certifique-se de atribuir à sua tabela final o nome `above_eight`!


```python
above_eight = ...
above_eight
```

```python
grader.check("q41")
```

**Questão 4.2.** Use `num_rows` (e aritmética) para encontrar a *proporção* de filmes no conjunto de dados que foram lançados entre 1900-1999, e a *proporção* de filmes no conjunto de dados que foram lançados no ano 2000 ou posterior.

Atribua `proportion_in_20th_century` à proporção de filmes no conjunto de dados que foram lançados entre 1900-1999, e `proportion_in_21st_century` à proporção de filmes no conjunto de dados que foram lançados no ano 2000 ou posterior.

*Dica:* A *proporção* de filmes lançados nos anos 1900 é o *número* de filmes lançados nos anos 1900, dividido pelo *número total* de filmes.


```python
num_movies_in_dataset = ...
num_in_20th_century = ...
num_in_21st_century = ...
proportion_in_20th_century = ...
proportion_in_21st_century = ...
print("Proportion in 20th century:", proportion_in_20th_century)
print("Proportion in 21st century:", proportion_in_21st_century)
```

```python
grader.check("q42")
```

## 5. Resumo

Para sua referência, aqui está uma tabela de todas as funções e métodos que vimos neste laboratório. Aprenderemos mais métodos para adicionar a esta tabela na próxima semana!

|Nome|Exemplo|Propósito|
|-|-|-|
|`sort`|`tbl.sort("N")`|Cria uma cópia de uma tabela ordenada pelos valores em uma coluna|
|`where`|`tbl.where("N", are.above(2))`|Cria uma cópia de uma tabela apenas com as linhas que correspondem a algum *predicado*|
|`num_rows`|`tbl.num_rows`|Calcula o número de linhas em uma tabela|
|`num_columns`|`tbl.num_columns`|Calcula o número de colunas em uma tabela|
|`select`|`tbl.select("N")`|Cria uma cópia de uma tabela apenas com algumas das colunas|
|`drop`|`tbl.drop("N")`|Cria uma cópia de uma tabela sem algumas das colunas|

## 6. Submissão

Você terminou o Laboratório 2! Para verificar novamente seu trabalho, a célula abaixo executará novamente todos os testes do corretor automático.

**Passos importantes para a submissão:** 
1. Execute os testes e verifique se todos passaram.
2. Escolha **Salvar Notebook** no menu **Arquivo**, então **execute a última célula**. 
3. Clique no link para baixar o arquivo zip.
4. Em seguida, envie o arquivo zip para a tarefa correspondente de acordo com as instruções do seu professor. 

**É sua responsabilidade garantir que seu trabalho esteja salvo antes de executar a última célula.**

Shah e Katsu queriam parabenizá-lo por terminar o laboratório 2! 

![Katsu_Shah](shah_katsu.png)

## Submissão

Certifique-se de ter executado todas as células em seu notebook em ordem antes de executar a célula abaixo, para que todas as imagens/gráficos apareçam na saída. A célula abaixo irá gerar um arquivo zip para você enviar. **Por favor, salve antes de exportar!**

```python
# Primeiro salve seu notebook, depois execute esta célula para exportar sua submissão.

grader.export(pdf=False, run_tests=True)
```

Tests for q11 (points: None)
```
>>> new_year
2022
```


Tests for q21 (points: None)
```
>>> round(near_twenty, 8)
19.99909998
```


Tests for q211 (points: None)
```
>>> import math
>>> floor_of_pi == 3
True
```


Tests for q32 (points: None)
```
>>> num_farmers_markets_columns == 59
True
```


Tests for q33 (points: None)
```
>>> sorted(farmers_markets_locations.labels)
['MarketName', 'State', 'city', 'x', 'y']
```
```
>>> farmers_markets_locations.num_rows == 8546
True
```


Tests for q34 (points: None)
```
>>> farmers_markets_without_fmid.num_columns
57
```
```
>>> print(sorted(farmers_markets_without_fmid.labels))
['Bakedgoods', 'Beans', 'Cheese', 'Coffee', 'County', 'Crafts', 'Credit', 'Eggs', 'Facebook', 'Flowers', 'Fruits', 'Grains', 'Herbs', 'Honey', 'Jams', 'Juices', 'Location', 'Maple', 'MarketName', 'Meat', 'Mushrooms', 'Nursery', 'Nuts', 'Organic', 'OtherMedia', 'PetFood', 'Plants', 'Poultry', 'Prepared', 'SFMNP', 'SNAP', 'Seafood', 'Season1Date', 'Season1Time', 'Season2Date', 'Season2Time', 'Season3Date', 'Season3Time', 'Season4Date', 'Season4Time', 'Soap', 'State', 'Tofu', 'Trees', 'Twitter', 'Vegetables', 'WIC', 'WICcash', 'Website', 'WildHarvested', 'Wine', 'Youtube', 'city', 'street', 'x', 'y', 'zip']
```


Tests for q35 (points: None)
```
>>> type(farmers_markets_locations_by_latitude) == tables.Table
True
```
```
>>> # HINT: Check the order of your table. 
>>> list(np.round(list(farmers_markets_locations_by_latitude.column('y').take(range(3))),6)) == [64.86275, 64.8459, 64.844414]
True
```


Tests for q36 (points: None)
```
>>> berkeley_markets.num_rows == 3
True
```
```
>>> list(berkeley_markets.column('city')) == ['Berkeley', 'Berkeley', 'Berkeley']
True
```


Tests for q41 (points: None)
```
>>> type(above_eight) == tables.Table
True
```
```
>>> above_eight.num_rows == 20
True
```
```
>>> # Make sure your columns are in the right order!
>>> # First column should be 'Title', second column should be 'Rating'
>>> print(above_eight.sort(0).take([17]))
Title       | Rating
Toy Story 3 | 8.3
```


Tests for q42 (points: None)
```
>>> proportion_in_20th_century == 0.684
True
```
```
>>> proportion_in_21st_century == 0.316
True
```
