!pip install otter

```python
# Inicialize o Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

# Trabalho de Casa 4: Funções, Tabelas e Grupos

Por favor, complete este notebook preenchendo as células fornecidas. Antes de começar, execute a célula anterior para carregar os testes fornecidos.

**Recurso útil:**
- [Referência Python](http://data8.org/sp22/python-reference.html): Folha de dicas de métodos úteis de array e tabela usados em Data 8!

**Leituras recomendadas**: 

* [Visualizando Distribuições Numéricas](https://www.inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html)
* [Funções e Tabelas](https://www.inferentialthinking.com/chapters/08/Functions_and_Tables.html)

Por favor, complete este notebook preenchendo as células fornecidas. Antes de começar, execute a célula a seguir para configurar o notebook importando algumas bibliotecas úteis. Toda vez que você iniciar seu servidor, precisará executar esta célula novamente.

Para todos os problemas que você deve escrever explicações e frases, você **deve** fornecer sua resposta no espaço designado. **Além disso, durante este trabalho de casa e todos os futuros, certifique-se de não reatribuir variáveis ao longo do notebook!** Por exemplo, se você usar `max_temperature` em sua resposta para uma pergunta, não a reatribua mais tarde. Caso contrário, você falhará em testes que pensou que estava passando anteriormente!

**Nota: Este trabalho de casa tem testes ocultos. Isso significa que mesmo que os testes digam que 100% passaram, isso não significa que sua nota final será 100%. Estaremos executando mais testes de correção depois que todos entregarem o trabalho de casa.**

Compartilhar respostas diretamente não é permitido, mas discutir problemas com a equipe do curso ou com outros alunos é incentivado.

Você deve começar cedo para ter tempo de obter ajuda se estiver preso.

## 1. San Diego cheio de Burritos

```python
# Execute esta célula para configurar o notebook, mas por favor não a altere.

# Estas linhas importam os módulos Numpy e Datascience.
import numpy as np
from datascience import *
import d8error

# Estas linhas fazem algumas mágicas de plotagem.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

import warnings
warnings.simplefilter('ignore', FutureWarning)
warnings.filterwarnings("ignore")
```

Mira, Sofia e Sara estão tentando usar a Ciência de Dados para encontrar os melhores burritos em San Diego! Suas amigas Jessica e Sonya forneceram-lhes dois conjuntos de dados abrangentes sobre muitos estabelecimentos de burritos na área de San Diego, obtidos e limpos de: https://www.kaggle.com/srcole/burritos-in-san-diego/data

A célula a seguir lê uma tabela chamada `ratings` que contém nomes de restaurantes de burritos, sua classificação no Yelp, classificação no Google, bem como sua classificação geral. A classificação `Overall` não é uma média das classificações `Yelp` e `Google`, mas sim a classificação geral dos clientes que foram pesquisados no estudo acima.

Também lê uma tabela chamada `burritos_types` que contém nomes de restaurantes de burritos, seus itens de menu, e o custo do respectivo item de menu no restaurante.

```python
# Apenas execute esta célula
ratings = Table.read_table("ratings.csv")
ratings.show(5)
burritos_types = Table.read_table("burritos_types.csv").drop(0)
burritos_types.show(5)
```

**Questão 1.** Seria mais fácil se pudéssemos combinar as informações em ambas as tabelas. Atribua `burritos` ao resultado da junção das duas tabelas, para que tenhamos uma tabela com as classificações para cada item correspondente do menu de cada restaurante. Cada item do menu tem a mesma classificação do restaurante de onde é originário. **(8 Pontos)**

*Nota:* Não importa qual tabela você coloca como argumento para o método da tabela, qualquer ordem funcionará para os testes do autograder.

*Dica:* Se você precisa de ajuda para usar o método `join`, consulte a [Folha de Referência Python](http://data8.org/sp22/python-reference.html) ou a [Seção 8.4](https://www.inferentialthinking.com/chapters/08/4/Joining_Tables_by_Columns.html) no livro didático.

```python
burritos = ...
burritos.show(5)
```

```python
grader.check("q1_1")
```

<!-- BEGIN QUESTION -->

**Questão 2.** Vamos ver como as pontuações do Yelp se comparam às pontuações do Google na tabela `burritos`. Primeiro, atribua `yelp_and_google` a uma tabela contendo apenas as colunas `Yelp` e `Google`. Em seguida, faça um gráfico de dispersão com as pontuações do Yelp no eixo x e as pontuações do Google no eixo y. **(8 Pontos)**


```python
yelp_and_google = ...
...

# Não altere/edite/remova a linha seguinte.
# Para ajudá-lo a tirar conclusões, traçamos uma linha reta no gráfico (y=x).
plt.plot(np.arange(2.5,5,.5), np.arange(2.5,5,.5));
```

```python
grader.check("q1_2")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Questão 3.** Olhando para o gráfico de dispersão que você acabou de fazer na Questão 1.2, você percebe algum padrão(s) (ou seja, um dos dois tipos de pontuações é consistentemente maior que o outro)? Se sim, descreva-os **brevemente** na célula abaixo. **(8 Pontos)**


_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

Aqui está uma revisão de como `.group` funciona! Você pode ler como `.group` funciona no [livro didático](https://www.inferentialthinking.com/chapters/08/2/Classifying_by_One_Variable.html), ou você pode assistir ao vídeo abaixo. O recurso em vídeo foi feito por um ex-membro da equipe, Divyesh Chotai!

Você também pode usar o [Visualizador de Funções de Tabela](http://data8.org/interactive_table_functions/) para obter mais experiência prática com a função `.group`.

```python
from IPython.display import YouTubeVideo
YouTubeVideo("HLoYTCUP0fc")
```

**Questão 4.** Há tantos tipos de burritos da Califórnia na tabela `burritos`! Sara quer saber qual tipo é o mais bem avaliado em todos os restaurantes. Para efeito destas questões, estamos tratando a classificação de cada item do menu como a do respectivo restaurante, já que não temos a classificação de todos os itens individuais nestes restaurantes. Você não precisa se preocupar com este fato, mas achamos importante mencioná-lo!

Crie uma tabela com duas colunas: a primeira coluna deve incluir os nomes dos burritos e a segunda coluna deve conter a classificação geral média desse burrito em todos os restaurantes. **Em seus cálculos, você deve comparar apenas burritos que contêm a palavra "California".** Por exemplo, existem burritos "California", burritos "California Breakfast", burritos "California Surf And Turf", etc. **(9 Pontos)**

*Dica:* Se vários restaurantes servem o burrito "California - Chicken", que método de tabela podemos usar para agregar esses juntos e encontrar a classificação geral média?

*Nota:* Para referência, a solução da equipe usou apenas uma linha. No entanto, sinta-se à vontade para dividir a solução em várias linhas e etapas; apenas certifique-se de atribuir a tabela final de saída a `california_burritos`!

```python
california_burritos = ...
california_burritos
```

```python
grader.check("q1_4")
```

**Questão 5.** Dada esta nova tabela `california_burritos`, Sara pode descobrir o nome do burrito da Califórnia com a classificação média geral mais alta! Atribua `best_california_burrito` a uma linha de código que produza a string que representa o nome do burrito da Califórnia com a classificação média geral mais alta. Se vários burritos satisfizerem este critério, você pode exibir qualquer um deles. **(8 Pontos)**

```python
best_california_burrito = ...
best_california_burrito
```

```python
grader.check("q1_5")
```

<!-- BEGIN QUESTION -->

**Questão 6.** Mira acredita que os burritos em San Diego são mais baratos (e têm melhor sabor) do que os burritos em Berkeley. Plote um histograma que visualize a distribuição dos custos dos burritos de San Diego na tabela `burritos`. Também use a variável `bins` fornecida ao fazer seu histograma, para que o histograma seja mais visualmente informativo. **(8 Pontos)**


```python
bins = np.arange(0, 15, 1)
# Por favor, também use os bins fornecidos
...
```

<!-- END QUESTION -->

**Questão 7.** Qual porcentagem de burritos em San Diego custa menos de $6? Atribua `burritos_less_than_6` à sua resposta, que deve estar entre 0 e 100. Você deve usar apenas o histograma acima para responder à pergunta. Não use código na tabela para encontrar a resposta, apenas observe as alturas e use Python para avaliar sua aritmética! **(8 Pontos)**

*Nota*: Sua resposta não precisa ser exata, mas deve estar dentro de algumas porcentagens da resposta da equipe.


```python
burritos_less_than_6 = ...
```

```python
grader.check("q1_7")
```

## 2. Salários dos Funcionários da Cidade de São Francisco

Este exercício é projetado para dar a você prática com o uso dos métodos `.pivot` e `.group` da Tabela. Aqui está um link para a [Folha de Referência Python](http://data8.org/sp22/python-reference.html) caso você precise de uma rápida revisão. O [Visualizador de Funções de Tabela](http://data8.org/interactive_table_functions/) também pode ser uma ferramenta útil.

Execute a célula abaixo para ver uma demonstração de como você pode usar pivot em uma tabela. (Obrigado ao ex-membro da equipe Divyesh Chotai!)

```python
from IPython.display import YouTubeVideo
YouTubeVideo("4WzXo8eKLAg")
```

A fonte de dados que usaremos nesta parte do trabalho de casa é [publicamente fornecida](https://data.sfgov.org/City-Management-and-Ethics/Employee-Compensation/88g8-5mnd/data) pela Cidade de São Francisco. Filtramos para reter apenas as colunas relevantes e restringimos os dados ao ano civil de 2019. Execute a célula a seguir para carregar nossos dados em uma tabela chamada `full_sf`.

```python
full_sf = Table.read_table("sf2019.csv")
full_sf.show(5)
```

A tabela tem uma linha para cada um dos 44.525 funcionários do governo de São Francisco em 2019.

As quatro primeiras colunas descrevem o trabalho do funcionário. Por exemplo, o funcionário na terceira linha da tabela tinha um trabalho chamado "IS Business Analyst-Senior". Chamaremos isso de *posição* ou *cargo* do funcionário. O trabalho estava em uma Família de Trabalho chamada Sistemas de Informação (daí o IS no título do trabalho), e estava no Departamento de Liberdade Condicional para Adultos que faz parte do Grupo de Organização de Proteção Pública do governo. Você estará trabalhando principalmente com a coluna `Job`.

As próximas três colunas contêm os valores em dólares pagos ao funcionário no ano civil de 2019 para salário, horas extras e benefícios. Note que o salário de um funcionário não inclui seus ganhos de horas extras.

A última coluna contém a remuneração total paga ao funcionário. É a soma das três colunas anteriores:

$$\text{Remuneração Total} = \text{Salário} + \text{Horas Extras} + \text{Benefícios}$$

Para este trabalho de casa, estaremos usando as seguintes colunas:
1. `Organization Group`: Um grupo de departamentos. Por exemplo, o Grupo Org. **Proteção Pública** inclui departamentos como Polícia, Bombeiros, Proteção para Adultos, Promotoria, etc.
2. `Department`: A unidade organizacional primária usada pela Cidade e Condado de São Francisco.
3. `Job`: A posição específica que um determinado trabalhador preenche.
4. `Total Compensation`: A soma do salário, horas extras e benefícios de um trabalhador em 2019.

Execute a célula a seguir para selecionar as colunas relevantes e criar uma nova tabela chamada `sf`.

```python
sf = full_sf.select("Job", "Department", "Organization Group",  "Total Compensation")
sf.show(5)
```

Queremos usar esta tabela para gerar arrays com os títulos dos cargos dos membros de cada **Grupo de Organização**.

**Questão 1.** Defina `job_titles` como uma tabela com duas colunas. A primeira coluna deve ser chamada `Organization Group` e ter o nome de cada "Organization Group" uma vez, e a segunda coluna deve ser chamada `Jobs` com cada linha nessa segunda coluna contendo um *array* com os nomes de todos os títulos de trabalho dentro desse "Organization Group". Não se preocupe se houver vários dos mesmos títulos de trabalho. **(9 Pontos)**

*Dica:* Pense em como `group` funciona: ele coleta valores em um array e então aplica uma função a esse array. Definimos duas funções abaixo para você, e você precisará usar uma delas em sua chamada para `group`.

```python
# Escolha uma das duas funções definidas abaixo em sua chamada para agrupar.
def first_item(array):
    '''Retorna o primeiro item'''
    return array.item(0)

def full_array(array):
    '''Retorna o array que é passado'''
    return array 

# Faça uma chamada para agrupar usando uma das funções acima quando você definir job_titles
job_titles = ...
job_titles
```

```python
grader.check("q2_1")
```

#### Entender o código que você acabou de escrever em 2.1 é importante para avançar na aula! Se você fez um palpite sortudo, reserve um tempo para olhar o código, passo a passo. O horário de atendimento é sempre um ótimo recurso!

**Questão 2.** No momento, a coluna `Job` da tabela `sf` não está ordenada (sem ordem específica). Os arrays que você gerou na coluna `Jobs` da questão anterior seriam os mesmos se tivéssemos ordenado alfabeticamente antes de gerá-los? Explique sua resposta. Para receber a pontuação total, sua resposta deve referenciar *como* o método `.group` funciona, e como ordenar a coluna `Jobs` afetaria isso. **(8 Pontos)**

*Nota:* Dois arrays são **iguais** se contêm o mesmo número de elementos e os elementos localizados nos índices correspondentes nos dois arrays são idênticos. Um exemplo de arrays que NÃO são iguais: `array([1,2]) != array([2,1])`.

_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

**Questão 3.** Defina `department_ranges` como uma tabela contendo departamentos como as linhas, e os grupos de organização como as colunas. Os valores nas linhas devem corresponder a um intervalo de remuneração total, onde o intervalo é definido como a **diferença entre a remuneração total mais alta e a remuneração total mais baixa no departamento para aquele grupo de organização**. **(9 Pontos)**

*Dica 1:* Primeiro, você precisará definir uma nova função `compensation_range` que recebe um array de remunerações e retorna o intervalo de remunerações nesse array.

*Dica 2:* Que função de tabela permite especificar as linhas e colunas de uma nova tabela? Você provavelmente assistiu a um vídeo sobre isso anteriormente no trabalho de casa!


```python
# Primeiro defina compensation_range
...
    ...

department_ranges = ...
department_ranges
```

```python
grader.check("q2_3")
```

<!-- BEGIN QUESTION -->

**Questão 4.** Dê uma explicação de por que alguns dos valores das linhas são `0` na tabela `department_ranges` da questão anterior. **(8 Pontos)**


_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

**Questão 5.** Encontre o número de departamentos que aparecem na tabela `sf` que têm uma remuneração total média superior a 125.000 dólares; atribua este valor à variável `num_over_125k`. **(9 Pontos)**

*Dica:* Os nomes das variáveis fornecidas são destinados a ajudar nos passos intermediários e no processo geral de pensamento. Sinta-se à vontade para excluí-los se preferir começar do zero, mas certifique-se de que sua resposta final seja atribuída a `num_over_125k`!

```python
depts_and_comp = ...
avg_of_depts = ...
num_over_125k = ...
num_over_125k
```

```python
grader.check("q2_5")
```

Você terminou o Trabalho de Casa 4!

**Passos importantes para a submissão:**
1. Execute os testes e verifique se todos passaram.
2. Escolha **Salvar Notebook** no menu **Arquivo**, então **execute a última célula**.
3. Clique no link para baixar o arquivo zip.
4. Em seguida, envie o arquivo zip para a tarefa correspondente de acordo com as instruções do seu instrutor.

**É sua responsabilidade garantir que seu trabalho esteja salvo antes de executar a última célula.**

## Submissão

Certifique-se de ter executado todas as células em seu notebook em ordem antes de executar a célula abaixo, para que todas as imagens/gráficos apareçam na saída. A célula abaixo irá gerar um arquivo zip para você enviar. **Por favor, salve antes de exportar!**

```python
# Primeiro salve seu notebook, depois execute esta célula para exportar sua submissão.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 0])
```
>>> # Make sure you have all the columns from both tables
>>> set(["Name", "Menu_Item", "Yelp", "Google", "Overall", "Cost"]) == set(burritos.labels)
True
```
```
>>> burritos.num_rows == 212
True
```


Tests for q1_2 (points: None)
```
>>> # Make sure you are using the correct table! 
>>> yelp_and_google.num_rows == 212
True
```


Tests for q1_4 (points: [1, 1])
```
>>> # Number of columns should be 2
>>> california_burritos.num_columns == 2
True
```
```
>>> # Number of rows should be 19
>>> california_burritos.num_rows == 19
True
```


Tests for q1_5 (points: [0])
```
>>> # Your answer should be a string
>>> type(best_california_burrito) == str
True
```


Tests for q1_7 (points: [0])
```
>>> # Your answer should be between 0 and 100.
>>> 0 <= burritos_less_than_6 <= 100
True
```


Tests for q2_1 (points: [0, 0, 0, 0])
```
>>> job_titles.num_columns
2
```
```
>>> job_titles.num_rows
6
```
```
>>> # Make sure that you have the correct column labels!
>>> np.asarray(job_titles.labels).item(1) != "Job full_array"
True
```
```
>>> # Make sure that you have the correct column labels!
>>> np.asarray(job_titles.labels).item(1) == "Jobs"
True
```


Tests for q2_3 (points: [0, 0, 2])
```
>>> # Double check that your salary_range function is correct
>>> compensation_range(make_array(5, 1, 20, 1000)) == 999
True
```
```
>>> # Make sure that the table has departments as the rows and organization groups as the columns.
>>> set(["Department", "Community Health", "Culture & Recreation", "General Administration & Finance", "Human Welfare & Neighborhood Development", "Public Protection", "Public Works, Transportation & Commerce"]) == set(department_ranges.labels)
True
```
```
>>> sum(department_ranges.column(1))
554179.0
```


Tests for q2_5 (points: [2, 2])
```
>>> # Make sure that your final answer is a number
>>> isinstance(num_over_125k, int)
True
```
```
>>> # Make sure that your answer makes sense given the sf table
>>> 0 <= num_over_125k <= 51
True
```
