```python
# Inicialize o Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

# Laboratório 8: Distribuição Normal e Variabilidade das Médias Amostrais

Bem-vindo ao Laboratório 8!

No laboratório de hoje, aprenderemos sobre [a variabilidade das médias amostrais](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) bem como [a distribuição normal](https://www.inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html).

**Submissão**: Quando terminar, execute todas as células, exceto a última, selecione Arquivo > Salvar Notebook e, em seguida, execute a célula final. Em seguida, envie o arquivo zip baixado, que inclui seu notebook, de acordo com as instruções do seu professor.

Primeiro, configure o notebook executando a célula abaixo.

```python
# Execute esta célula para configurar o notebook, mas por favor não a altere.

# Estas linhas importam os módulos Numpy e Datascience.
import numpy as np
from datascience import *

# Estas linhas fazem algumas mágicas de plotagem.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')

import d8error
```

# 1. Distribuições Normais

Quando visualizamos a distribuição de uma amostra, geralmente estamos interessados na média e no desvio padrão da amostra (pelo resto deste laboratório, abreviaremos "desvio padrão" como "DP"). Essas duas estatísticas de amostra podem nos dar uma visão geral da distribuição - informando-nos onde a distribuição se situa na linha numérica e quão dispersa ela está, respectivamente.

Queremos verificar se os dados estão relacionados linearmente, então devemos olhar para os dados.

**Questão 1.1.** A próxima célula carrega a tabela `births` da aula, que é uma grande amostra aleatória de nascimentos nos EUA e inclui informações sobre pares mãe-filho.

Trace a distribuição das idades das mães a partir da tabela. Não altere a última linha, que irá plotar a média da amostra na própria distribuição.

```python
births = Table.read_table('baby.csv')
...

# Não altere nada abaixo desta linha.
plots.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

Pelo gráfico acima, podemos ver que a média é o centro de gravidade ou ponto de equilíbrio da distribuição. Se você recortar a distribuição de um pedaço de papelão e, em seguida, colocar o dedo na média, a distribuição se equilibraria perfeitamente no seu dedo. Como a distribuição acima é assimétrica à direita (o que significa que tem uma longa cauda à direita), sabemos que a média da distribuição é maior que a mediana, que é o ponto "intermediário" dos dados. Por outro lado, se a distribuição fosse assimétrica à esquerda, saberíamos que a média seria menor que a mediana.

**Questão 1.2.** Execute a célula a seguir para comparar a média (vermelho) e a mediana (azul) da distribuição das idades das mães.

```python
births.hist("Maternal Age")
plots.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50, label="Mean")
plots.scatter(np.median(births.column("Maternal Age")), 0, color='blue', zorder=2, s=50, label="Median")
plots.legend()
plots.ylim(-0.005, 0.07);
```

Também estamos interessados no desvio padrão das idades das mães. O DP nos dá uma noção de quão variáveis são as idades das mães em torno da idade média das mães. Se o DP é grande, então as idades das mães devem se espalhar por uma grande faixa a partir da média. Se o DP é pequeno, então as idades das mães devem estar fortemente agrupadas em torno da idade média da mãe.

**O DP de um array é definido como a raiz quadrada média dos desvios, onde os desvios são a diferença entre os elementos originais e sua média.**

Você pode ler mais sobre Variabilidade e DP no [Capítulo 14.2](https://inferentialthinking.com/chapters/14/2/Variability.html). Curiosidade! $\sigma$ (letra grega sigma) é usada para representar o DP e $\mu$ (letra grega mu) é usada para a média.

**Question 1.3.** Assign `age_mean` to the mean of the maternal ages and `age_sd` to the SD of the maternal ages. Then run the cell below to see the width of one SD (blue) from the sample mean (red) plotted on the histogram of maternal ages.


```python
age_mean = ...
age_sd = ...
births.hist("Maternal Age")

plots.scatter(age_mean, 0, color='red', zorder=2, s=50);
plots.scatter(age_mean+age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(age_mean-age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

```python
grader.check("q1_3")
```

Neste histograma, o desvio padrão não é fácil de identificar apenas olhando para o gráfico.

No entanto, as distribuições de algumas variáveis nos permitem identificar facilmente o desvio padrão no gráfico. Por exemplo, se uma amostra segue uma *distribuição normal*, o desvio padrão é facilmente identificado no ponto de inflexão (o ponto onde a curva começa a mudar a direção de sua curvatura) da distribuição.

**Questão 1.4.** Preencha o seguinte código para examinar a distribuição das alturas maternas, que é aproximadamente normalmente distribuída. Vamos plotar o desvio padrão no histograma, como antes - observe onde um desvio padrão (azul) longe da média (vermelho) cai no gráfico.

```python
height_mean = ...
height_sd = ...
births.hist("Maternal Height", bins=np.arange(55,75,1))

plots.scatter((height_mean), 0, color='red', zorder=2, s=50);
plots.scatter(height_mean+height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(height_mean-height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.01, 0.16);
```

```python
grader.check("q1_4")
```

Nem sempre sabemos como uma variável será distribuída, e fazer suposições sobre se uma variável seguirá ou não uma distribuição normal é perigoso. No entanto, o **Teorema do Limite Central** define uma distribuição que sempre segue uma distribuição normal. **A distribuição das *somas* e *médias* de todas as grandes amostras aleatórias retiradas com reposição de uma única distribuição (independentemente da forma original da distribuição) será normalmente distribuída.** Lembre-se de que o Teorema do Limite Central se refere à distribuição de uma *estatística* calculada a partir de uma distribuição, não à distribuição da amostra original ou da população. Se isso for confuso, pergunte a um membro da equipe!

A próxima seção explorará as distribuições das médias amostrais, e você verá como o desvio padrão dessas distribuições depende dos tamanhos das amostras.

# 2. Variabilidade da Média Amostral

Pelo [Teorema do Limite Central](https://www.inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html), a distribuição de probabilidade da média ou soma de uma grande amostra aleatória é aproximadamente normal. A curva em forma de sino é centrada na média da população. Algumas das médias amostrais são maiores e algumas são menores, mas os desvios da média da população são aproximadamente simétricos de ambos os lados, como vimos repetidamente.

Em nossas simulações, também notamos que as médias de amostras maiores tendem a ser mais agrupadas em torno da média da população do que as médias de amostras menores. Nesta seção, quantificaremos a [variabilidade da média amostral](https://www.inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) e desenvolveremos uma relação entre a variabilidade e o tamanho da amostra.

Vamos dar uma olhada nos salários dos funcionários da cidade de São Francisco em 2014. O salário médio relatado pelo governo da cidade era de cerca de $75,463.92.

*Nota: Se você ficar preso em qualquer parte deste laboratório, por favor, consulte o [Capítulo 14 do livro didático](https://www.inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html).*

```python
salaries = Table.read_table('sf_salaries_2014.csv').select("salary")
salaries.show(5)
```

```python
salary_mean = np.mean(salaries.column('salary'))
print('Mean salary of San Francisco city employees in 2014: ', salary_mean.round(2))
```

```python
salaries.hist('salary', bins=np.arange(0, 250000+10000*2, 10000))
plots.scatter(salary_mean, 0, color='red', zorder=2, s=50)
plots.title('2014 salaries of city of SF employees')
plots.ylim(-0.0000005, 0.000014);
```

Claramente, a população *não* segue uma distribuição normal. Tenha isso em mente à medida que avançamos nestes exercícios.

Vamos pegar amostras aleatórias **com reposição** e olhar para a distribuição de probabilidade da média da amostra. Como de costume, usaremos simulação para obter uma aproximação empírica desta distribuição.

**Questão 2.1.** Defina uma função `one_sample_mean`. Seus argumentos devem ser `table` (o nome de uma tabela), `label` (o rótulo da coluna contendo a variável) e `sample_size` (o número de funcionários na amostra). Ela deve amostrar com reposição da tabela e retornar a média da coluna `label` da amostra.

```python
def one_sample_mean(table, label, sample_size):
    new_sample = ...
    new_sample_mean = ...
    ...
```

```python
grader.check("q2_1")
```

**Questão 2.2.** Use `one_sample_mean` para definir uma função `simulate_sample_mean`. Os argumentos são o nome da `table`, o `label` da coluna contendo a variável, o `sample_size`, e o número de simulações (`repetitions`).

A função deve amostrar com reposição da tabela e calcular a média de cada amostra. Ela deve salvar as médias da amostra em um array chamado `means`. O código restante na função exibe um histograma empírico das médias da amostra.

*Nota: Arredondamos todos os valores para 2 casas decimais porque estamos trabalhando com salários.*

```python
"""Empirical distribution of random sample means"""

def simulate_sample_mean(table, label, sample_size, repetitions):
    
    means = make_array()

    for i in np.arange(repetitions):
        new_sample_mean = ...
        means = ...

    sample_means = Table().with_column('Sample Means', means)
    
    # Exibir histograma empírico e imprimir todas as quantidades relevantes - não altere isso!
    sample_means.hist(bins=20)
    plots.xlabel('Sample Means')
    plots.title('Sample Size {sample_size}; {repetitions} Resamples'.format(sample_size=sample_size, repetitions=repetitions))
    print("Sample size: ", sample_size)
    print("Population mean:", np.mean(table.column(label)).round(2))
    print("Average of sample means: ", np.mean(means).round(2))
    print("Population SD:", np.std(table.column(label)).round(2))
    print("SD of sample means:", np.std(means).round(2))
    return np.std(means)
```

Verifique com um colega ou TA se você implementou a função acima corretamente. Se você não a implementou corretamente, o restante do laboratório não funcionará adequadamente, então esta etapa é crucial.

Na célula a seguir, criaremos uma amostra de tamanho 100 de `salaries` e a representaremos graficamente usando nossa nova função `simulate_sample_mean`.

*Dica: Você deve ver uma distribuição semelhante a algo que temos discutido. Se não, verifique sua função.*

```python
simulate_sample_mean(salaries, 'salary', 100, 10000) 
plots.xlim(50000, 100000);
```

Observe que nossa distribuição de *médias amostrais* parece aproximadamente normal! Importa se nossa distribuição original era normalmente distribuída?

**Questão 2.3.** Simule duas médias de amostras, uma para uma amostra de 400 salários e outra para uma amostra de 625 salários. Em cada caso, realize 10.000 repetições. Não se preocupe com as linhas `plots.xlim` ou `plots.ylim` - elas apenas garantem que todos os gráficos tenham o mesmo eixo x e eixo y, respectivamente.

```python
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
print('\n')
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
```

**Questão 2.4.** Atribua `q2_4` a um array de números correspondentes às afirmações verdadeiras sobre os gráficos da questão 2.3.

*Dica*: Se uma distribuição tem forma de sino, que tipo de curva ela segue?

1. Vemos o Teorema do Limite Central (TLC) em ação porque as distribuições das médias amostrais têm forma de sino.
2. Vemos a Lei dos Grandes Números em ação porque as distribuições das médias amostrais se parecem com a distribuição da população.
3. Uma das condições para o TLC é que temos que tirar uma pequena amostra aleatória com reposição da população.
4. Uma das condições para o TLC é que temos que tirar uma grande amostra aleatória com reposição da população.
5. Uma das condições para o TLC é que a população deve ser normalmente distribuída.

```python
q2_4 = ...
```

```python
grader.check("q2_4")
```

**Questão 2.5.** Atribua `q2_5` a um array de números correspondentes às afirmações verdadeiras sobre os gráficos da questão 2.3.

1. Ambos os gráficos em 2.3 estão aproximadamente centrados na média da população.
2. Ambos os gráficos em 2.3 estão aproximadamente centrados na média de uma amostra específica.
3. A distribuição das médias amostrais para o tamanho da amostra 625 tem menos variabilidade do que a distribuição das médias amostrais para o tamanho da amostra 400.
4. A distribuição das médias amostrais para o tamanho da amostra 625 tem mais variabilidade do que a distribuição das médias amostrais para o tamanho da amostra 400.

```python
q2_5 = ...
```

```python
grader.check("q2_5")
```

A seguir, veremos o que acontece quando pegamos um **número crescente de reamostras de um tamanho de amostra fixo.** Observe qual número no código muda e o que permanece o mesmo. Como a distribuição das médias reamostradas muda?

```python
simulate_sample_mean(salaries, 'salary', 100, 500)
plots.xlim(50000, 100000)
plots.ylim(0, 0.00009);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 1000)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 5000)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 10000)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

O que você notou sobre as distribuições das médias amostrais nos quatro histogramas acima? Discuta com seus colegas ou pergunte a um membro da equipe.

**Pergunta 2.6.** Atribua à variável `SD_of_sample_means` o número inteiro correspondente à sua resposta para a seguinte pergunta:

Quando eu aumento o número de *re-amostras* que faço, para um tamanho de amostra *fixo*, o desvio padrão das médias das minhas amostras...

1. Aumentará
2. Diminuirá
3. Permanecerá aproximadamente o mesmo
4. Variará amplamente

```python
SD_of_sample_means = ...
```

```python
grader.check("q2_6")
```

**Questão 2.7.** Vamos pensar sobre como as relações entre o desvio padrão da população, o desvio padrão da amostra e o desvio padrão das médias das amostras mudam com o tamanho da amostra variável. Qual das seguintes afirmações é verdadeira? Atribua a variável `pop_vs_sample` a um array de número(s) inteiro(s) que corresponda(m) a afirmação(ões) verdadeira(s).

*Dica 1:* O desvio padrão da amostra é diferente do desvio padrão das médias das amostras.

1. O desvio padrão da amostra diminui com o aumento do tamanho da amostra.
2. O desvio padrão da amostra aumenta com o aumento do tamanho da amostra.
3. O desvio padrão da amostra se torna mais consistente com o desvio padrão da população com o aumento do tamanho da amostra.
4. O desvio padrão das médias das amostras diminui com o aumento do tamanho da amostra.
5. O desvio padrão das médias das amostras aumenta com o aumento do tamanho da amostra.
6. O desvio padrão das médias das amostras permanece o mesmo com o aumento do tamanho da amostra.

```python
pop_vs_sample = ...
```

```python
grader.check("q2_7")
```

**Questão 2.8.** Existe uma relação entre o tamanho da amostra e o desvio padrão das médias das amostras? Atribua `q2_8` ao número correspondente à afirmação que responde a esta pergunta.

*Dica:* O [Capítulo 14.5](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) do livro didático pode ser útil.

1. O desvio padrão das médias das amostras é inversamente proporcional à raiz quadrada do tamanho da amostra.
2. O desvio padrão das médias das amostras é diretamente proporcional à raiz quadrada do tamanho da amostra.

```python
q2_8 = ...
```

```python
grader.check("q2_8")
```

Ao longo deste laboratório, temos tirado muitas amostras aleatórias de uma população. No entanto, todos esses princípios são válidos para re-amostras bootstrap de uma única amostra. Se sua amostra original for relativamente grande, todas as suas re-amostras também serão relativamente grandes, e, assim, o desvio padrão das médias das re-amostras será relativamente pequeno.

Para mudar a variabilidade da média da sua amostra, você precisaria mudar o tamanho da amostra original da qual você está tirando as re-amostras bootstrap.

## 3. Submissão

<img src="jim.jpg" alt="drawing" width="300"/>

Jim realmente queria parabenizá-lo por completar o Laboratório 8!

**Passos importantes para a submissão:** 
1. Execute os testes e verifique se todos passaram.
2. Escolha **Salvar Notebook** no menu **Arquivo**, então **execute a última célula**. 
3. Clique no link para baixar o arquivo zip.
4. Em seguida, envie o arquivo zip para a tarefa correspondente de acordo com as instruções do seu professor. 

**É sua responsabilidade garantir que seu trabalho esteja salvo antes de executar a última célula.**

## Submissão

Certifique-se de ter executado todas as células em seu notebook em ordem antes de executar a célula abaixo, para que todas as imagens/gráficos apareçam na saída. A célula abaixo gerará um arquivo zip para você enviar. **Por favor, salve antes de exportar!**

```python
# Primeiro salve seu notebook, depois execute esta célula para exportar sua submissão.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_3 (points: None)
```
>>> round(age_mean, 3) == 27.228
True
```
```
>>> round(age_sd, 3) == 5.815
True
```


Tests for q1_4 (points: None)
```
>>> round(height_mean, 3) == 64.049
True
```
```
>>> round(height_sd, 3) == 2.525
True
```


Tests for q2_1 (points: None)
```
>>> np.random.seed(8)
>>> np.round(one_sample_mean(salaries, 'salary', 100), 3) == 76699.829
True
```


Tests for q2_4 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(set(np.array(q2_4).astype(int))) 
'3649655f87d3a44afa82f25eced24944'
```


Tests for q2_5 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(set(np.array(q2_5).astype(int)))
'b793dd27381f7b63cada88b749c0e023'
```


Tests for q2_6 (points: None)
```
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(SD_of_sample_means))
'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```


Tests for q2_7 (points: None)
```
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(set(np.array(pop_vs_sample).astype(int)))
'1d919a653870c713cf8bd17b9c9c65ee'
```


Tests for q2_8 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(q2_8))
'c4ca4238a0b923820dcc509a6f75849b'
```
