```python
# Inicialize o Otter
import otter
grader = otter.Notebook("lab10.ipynb")
```

# Laboratório 10: Probabilidade Condicional
Bem-vindo ao Laboratório 10!

Este laboratório é uma introdução às probabilidades condicionais.

O laboratório inclui uma visualização chamada *array de ícones*. Ela é destinada a ser uma parte instrucional do laboratório para ajudar a construir intuições sobre a probabilidade condicional. Essas visualizações não aparecem no livro didático e não aparecerão em nenhum exame.

**Submissão**: Uma vez que você tenha terminado, execute todas as células exceto a última, selecione Arquivo > Salvar Caderno, e então execute a célula final. Depois, submeta o arquivo zip baixado, que inclui seu caderno, de acordo com as instruções do seu instrutor.

*Este conteúdo é protegido e não pode ser compartilhado, carregado ou distribuído.*

```python
# Execute esta célula para configurar o caderno, mas por favor, não a altere.

# Estas linhas importam os módulos Numpy e Datascience.
import numpy as np
from datascience import *

# Estas linhas fazem alguma mágica de plotagem sofisticada.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')

# Esta linha carrega o código de visualização para este laboratório.
import visualizations

import d8error
```

# 1. Para que serve a probabilidade condicional?

Suponha que temos uma população conhecida, como todos os cães na Califórnia. Até agora, vimos três maneiras de *prever* algo sobre um indivíduo dessa população, dado um conhecimento incompleto sobre a identidade do indivíduo:

* Se não sabemos nada sobre o cão individual, poderíamos prever que sua velocidade é a *média* ou *mediana* de todas as velocidades na população.
* Se conhecemos a altura do cão, mas não sua velocidade, poderíamos usar *regressão linear* para prever sua velocidade a partir de sua altura. A previsão resultante ainda é imperfeita, mas pode ser mais precisa do que a média da população.
* Se conhecemos a raça, altura e idade do cão, poderíamos usar *classificação de vizinho mais próximo* para prever sua velocidade comparando-o a uma coleção de cães com velocidades conhecidas.

Também podemos calcular probabilidades condicionais para fazer previsões sobre indivíduos ou eventos. Esta técnica é diferente dos métodos anteriores que examinamos porque

1. nossa previsão para cada resultado é descrita por uma probabilidade, e 
2. cada probabilidade pode ser exatamente calculada a partir de suposições, ao invés de estimada a partir de dados.

# 2. Arrays de ícones
As partes 3 e 4 deste laboratório trabalham com um exemplo mais complexo sobre doenças, mas primeiro, vamos começar com um exemplo simples.

Imagine que você e Samantha estão jogando um jogo no qual você recebe uma bolinha e tem a tarefa de determinar a textura e o tamanho da bolinha. Você não sabe nada sobre a bolinha que recebeu, mas sabe que Samantha a retirou **uniformemente ao acaso** de uma sacola que continha as seguintes bolinhas:
* 4 bolinhas grandes e brilhantes,
* 1 bolinha grande e opaca,
* 6 bolinhas pequenas e brilhantes,
* 2 bolinhas pequenas e opacas.

**Questão 2.0.1.** Sabendo apenas o que lhe dissemos até agora, qual é a probabilidade de que a bolinha que você recebeu seja uma bolinha grande e brilhante?

```python
probability_large_shiny = ...
```

```python
grader.check("q2_0_1")
```

Aqui está uma tabela com essas bolinhas:

```python
marbles = Table.read_table("marbles.csv")
marbles.show()
```

Aqui estão as contagens de cada tipo de bolinha em uma tabela dinâmica.

```python
marbles.pivot('surface', 'size')
```

Aqui estão todas as diferentes combinações de superfície e tamanho, com a contagem para cada combinação de superfície e tamanho. Cada tipo de bolinha aparece em sua própria linha.

```python
marbles.group(['surface', 'size'])
```

Incluímos algum código para exibir algo chamado *array de ícones*. As funções na célula abaixo criam arrays de ícones a partir de vários tipos de tabelas. Não se preocupe em entender o código; apenas execute esta célula.

**NOTA:** Você pode ignorar os rótulos do eixo y. Apenas lembre-se de que cada caixa representa 1 bolinha.

```python
# Execute esta célula.

#######################################################################
# As funções que você realmente precisará usar estão aqui. Cada uma é
# uma maneira de criar um array de ícones a partir de uma tabela formatada de maneira diferente.
#######################################################################

def display_icon_array(table, groups, individuals_name):
    """
    Dada uma tabela e algumas colunas para agrupá-la, exibe um array de ícones
    dos grupos.
    
    groups deve ser um array de rótulos de colunas na tabela.
    
    individuals_name é o seu nome para as linhas individuais da tabela.
    Por exemplo, se estamos falando sobre uma população de pessoas,
    individuals_name deve ser "pessoas".
    
    Por exemplo:
    
    display_icon_array(marbles, ["surface", "size"], "marbles")
    """
    display_grouped_icon_array(table.group(groups), individuals_name)

def display_grouped_icon_array(grouped_data, individuals_name):
    """
    Dada uma tabela com contagens para dados agrupados por uma ou mais categorias,
    exibe um array de ícones dos grupos representados na tabela.
    
    grouped_data deve ser uma tabela de frequências ou contagens, como
    uma tabela criada ao chamar o método groups em alguma tabela.
    
    individuals_name é o seu nome para os membros individuais do
    conjunto de dados. Por exemplo, se estamos falando sobre uma população de
    pessoas, individuals_name deve ser "pessoas".
    
    Por exemplo:
    
    display_grouped_icon_array(marbles.group(["surface", "size"]), "marbles")
    """
    visualizations.display_combinations(grouped_data, individuals_name=individuals_name)

def display_crosstab_icon_array(crosstabulation, x_label, individuals_name):
    """
    Dada uma tabela de crosstabulação, exibe um array de ícones dos grupos
    representados na tabela.
    
    crosstabulation deve ser uma tabela de frequências ou contagens criada ao
    chamar pivot em alguma tabela.
    
    x_label deve ser o rótulo das categorias listadas como colunas (no
    "eixo x" quando a tabela de crosstabulação é impressa).
    
    individuals_name é o seu nome para os membros individuais do
    conjunto de dados. Por exemplo, se estamos falando sobre uma população de
    pessoas, individuals_name deve ser "pessoas".
    
    Por exemplo:
    
    display_crosstab_icon_array(marbles.pivot("surface", "size"), "surface", "marbles")
    """
    display_grouped_icon_array(visualizations.pivot_table_to_groups(crosstabulation, x_label), individuals_name)
```

Aqui está um array de ícones de todas as bolinhas, agrupadas por superfície e tamanho:

```python
# Execute esta célula.
display_grouped_icon_array(marbles.group(["surface", "size"]), "marble(s)")
```

Você deve imaginar que a bolinha que lhe foi dada foi selecionada por um sorteio aleatório desses 13 ícones.

A seguir está um array de ícones das bolinhas, agrupadas **apenas pela superfície (brilhante/opaca)**.

```python
display_grouped_icon_array(marbles.group("surface"), "marble(s)")
```

Sabendo nada mais sobre a bolinha, é igualmente provável que seja qualquer uma das bolinhas representadas acima; isso porque assumimos que a bolinha foi selecionada **uniformemente ao acaso** da sacola.

**Questão 2.0.2.** Qual é a probabilidade de que você tenha recebido uma bolinha brilhante? Calcule isso manualmente (usando Python para aritmética) olhando para o seu array de ícones.

```python
probability_shiny = ...
```

```python
grader.check("q2_0_2")
```

## 2.1. Probabilidade Condicional

Suponha que você ouça Samantha dizer que você recebeu uma bolinha grande. Isso de alguma forma muda a chance de sua bolinha ser brilhante? Vamos descobrir.

Volte ao array completo de ícones, exibido abaixo para sua conveniência.

```python
display_grouped_icon_array(marbles.group(["surface", "size"]), "marble(s)")
```

Na questão 2.0.2, assumimos que a bolinha que você recebeu tinha a mesma probabilidade de ser qualquer uma das bolinhas, simplesmente porque não sabíamos de nada melhor. É por isso que olhamos para todas as bolinhas para calcular a probabilidade de sua bolinha ser brilhante.

Mas assumindo que você recebeu uma bolinha grande, podemos eliminar algumas dessas possibilidades. Em particular, você não poderia ter recebido uma bolinha pequena brilhante ou uma bolinha pequena opaca.

Você ainda tem a mesma probabilidade de ter recebido qualquer uma das bolinhas restantes, porque você não sabe nenhuma outra informação. Então, aqui está um array de ícones dessas possibilidades restantes:

```python
# Apenas execute esta célula.
display_grouped_icon_array(marbles.where("size", "large").group("surface"), "large marble(s)")
```

**Questão 2.1.1.** Qual é a probabilidade de Samantha lhe dar uma bolinha brilhante, sabendo que ela lhe deu uma bolinha grande?

*Dica: Use o array de ícones.*

```python
probability_shiny_given_large = ...
```

```python
grader.check("q2_1_1")
```

Você deve ter descoberto que isso é diferente da probabilidade de Samantha lhe dar uma bolinha brilhante, sem nenhuma informação sobre o tamanho, que você calculou anteriormente. A distribuição de superfícies entre as bolinhas grandes é um pouco diferente da distribuição de superfícies entre todas as bolinhas.

```python
# Execute esta célula para exibir o array de ícones. Em seguida, preencha e responda à próxima questão.
display_grouped_icon_array(marbles.where("surface", "shiny").group("size"), "shiny marble(s)")
```

**Questão 2.1.2.** Suponha, em vez disso, que Samantha disse que lhe deu uma bolinha **brilhante** (oba!). Qual é a probabilidade de que a bolinha que lhe foi dada seja grande?

Execute a célula de código acima para exibir o array de ícones, depois atribua `probability_large_given_shiny` ao valor apropriado.

```python
probability_large_given_shiny = ...
```

```python
grader.check("q2_1_2")
```

**Questão 2.1.3.** Você pode responder às duas questões anteriores apenas olhando para o array completo de ícones? (Você pode executar a célula abaixo para vê-lo novamente.)

*Dica:* Confira o [Capítulo 18.2](https://inferentialthinking.com/chapters/18/2/Making_Decisions.html) para uma revisão sobre como usar a Regra de Bayes.

```python
# Apenas execute esta célula. A próxima célula é onde você deve escrever sua resposta.
display_grouped_icon_array(marbles.group(["surface", "size"]), "marble(s)")
```

Se você pode, como? Se não, por quê? Verifique com seus colegas de laboratório ou um membro da equipe para ver se você está no caminho certo.

_Digite sua resposta aqui, substituindo este texto._

# 3. Rastreamento de Câncer
Esperamos que os arrays de ícones da parte anterior tenham ajudado a construir intuição sobre por que as probabilidades condicionais podem ser úteis. Agora, vamos olhar para uma aplicação na vida real.

### Contexto
O teste médico é muito importante, especialmente para o câncer. Um rastreamento básico de câncer envolve procurar por câncer antes que a pessoa apresente quaisquer sintomas.

O rastreamento de câncer não é 100% confiável e pode ter erros. Há casos em que alguém sem câncer ou com um tumor benigno vê um resultado de teste positivo. Também há casos em que alguém com câncer pode receber um resultado negativo. (“Positivo” implica uma indicação de câncer neste contexto.) O primeiro caso, chamado de falso positivo, pode causar ansiedade para um paciente e levar a mais testes que podem ser arriscados e envolver radiação desnecessária. No segundo caso, chamado de falso negativo, uma pessoa não receberia as informações necessárias e o possível tratamento para sua condição.

A probabilidade condicional pode fornecer insights sobre a precisão dos testes iniciais de câncer. Por exemplo, você pode calcular a chance de uma pessoa ter câncer dado os resultados de um teste diagnóstico combinando informações de diferentes distribuições de probabilidade. Você verá que a chance de a Pessoa X ter câncer, mesmo se receber um resultado positivo, não é necessariamente 100%.

É importante notar que as probabilidades condicionais oferecem apenas uma linha de insight para entender um resultado de teste e que essa abordagem estatística não leva em conta fatores importantes como hereditariedade ou ambiente que afetam a saúde e diagnósticos de uma pessoa. O [Centros de Controle e Prevenção de Doenças](https://www.cdc.gov/cancer/dcpc/prevention/screening.htm) (CDC) recomenda a realização de rastreamentos preventivos de câncer, e deve-se sempre consultar um médico se estiver preocupado com sua saúde.

## 3.1. Estatísticas básicas sobre câncer

Nota: Todas as estatísticas a seguir são inventadas e não refletem necessariamente o estado real do mundo.

Suponha que, em um grupo representativo de 10.000 pessoas testadas para câncer ("representativo" no sentido de que as frequências de diferentes eventos são as mesmas que as frequências na população inteira):
1. 100 pessoas têm câncer.
2. Entre as 100 pessoas que têm câncer, 90 têm resultados positivos em um teste de câncer e 10 têm resultados negativos. (Então, 10 pessoas recebem resultados falsos negativos.)
3. As outras 9.900 pessoas não têm câncer.
4. Entre essas 9.900 pessoas, 198 têm resultados positivos em um teste de câncer e as outras 9.702 têm resultados negativos. (Então, 198 veem resultados "falsos positivos".)

Abaixo, geramos uma tabela com dados dessas 10.000 pessoas hipotéticas.

*Nota: Essas estatísticas são inventadas e não refletem necessariamente o estado real do mundo.*

```python
people = Table().with_columns(
    "status", ["cancer", "cancer", "no cancer", "no cancer"],
    "test status", ["positive", "negative", "positive", "negative"],
    "count", [90, 10, 198, 9702])
people
```

Uma maneira de visualizar este conjunto de dados é com uma tabela de contingência, que você já viu antes.

**Questão 3.1.1.** Usando a tabela `people` definida acima, crie uma tabela de contingência que se pareça com esta:

|status|negativo|positivo|
|-|-|-|
|câncer|||
|sem câncer||||

...com a **contagem** de cada grupo preenchida, de acordo com o que lhe dissemos acima. As contagens nas 4 caixas devem somar 10.000.

```python
cancer = ...
cancer
```

```python
grader.check("q3_1_1")
```

Aqui estão os dados das `pessoas` em um array de ícones.

```python
display_grouped_icon_array(people, "people who've taken a cancer test")
```

Agora, vamos pensar em como você pode usar esse tipo de informação quando a Pessoa X é testada para câncer.

Antes de saber qualquer informação sobre a Pessoa X, você poderia imaginar a Pessoa X como uma **amostra aleatória uniforme** de uma das 10.000 pessoas nesta população imaginária de pessoas que foram testadas.

Qual é a chance de que a Pessoa X tenha câncer, sem saber mais nada sobre ela? É $\frac{100}{10000}$, ou 1%. Podemos ver isso mais diretamente com este array de ícones:

```python
by_health = people.select(0, 2).group(0, sum).relabeled(1, 'count')
display_grouped_icon_array(by_health, "people who've taken a cancer test")
```

**Questão 3.1.2.** Qual é a chance de que a Pessoa X tenha um resultado de teste positivo, sem saber mais nada sobre ela?
Execute a próxima célula de código para exibir um array de ícones, depois atribua `probability_positive_test` a este valor.

```python
# Execute esta célula primeiro para exibir um array de ícones. Depois, preencha a probabilidade de um resultado de teste positivo na última linha.
by_test = people.select(1, 2).group(0, sum).relabeled(1, 'count')
display_grouped_icon_array(by_test, "people who've taken a cancer test")

# Agora preencha a probabilidade de um resultado de teste positivo
probability_positive_test = ...
```

## 3.2. Interpretando os resultados dos testes
Suponha que a Pessoa X tenha um resultado de teste positivo. Isso significa que agora você pode restringi-la a fazer parte de apenas um dos dois seguintes grupos:
1. As pessoas com câncer que têm um resultado de teste positivo.
2. As pessoas sem câncer que têm um resultado de teste positivo.

Aqui está um array de ícones para esses dois grupos:

```python
# Apenas execute esta célula.
display_grouped_icon_array(people.where("test status", are.equal_to("positive")).drop(1), "people who have a positive test result")
```

A *probabilidade condicional* de que a Pessoa X **tenha câncer dado o seu resultado positivo no teste** é a chance de que ela esteja no primeiro grupo (câncer), assumindo que ela tenha um resultado positivo no teste.

**Questão 3.2.1.** Observando o array de ícones acima, a probabilidade condicional de que a Pessoa X tenha câncer **dado o seu resultado positivo no teste** está mais próxima de:

1. $\displaystyle\frac{9}{10}$

2. $\displaystyle\frac{2}{3}$

3. $\displaystyle\frac{1}{2}$

4. $\displaystyle\frac{1}{3}$

5. $\displaystyle\frac{1}{100}$


Atribua `rough_prob_cancer_given_positive` a um inteiro correspondente à sua resposta.

```python
# Defina isso para 1, 2, 3, 4 ou 5, correspondendo à probabilidade correta.
rough_prob_cancer_given_positive = ...
```

```python
grader.check("q3_2_1")
```

**Questão 3.2.2.** Agora escreva código para calcular essa probabilidade exatamente, usando a tabela de contingência original que você escreveu (a tabela `cancer`).

Execute a próxima célula de código para ver a tabela `cancer`, depois preencha `prob_cancer_given_positive` com seu código.

```python
# Execute esta célula primeiro para ver a tabela de câncer. Depois, preencha a próxima linha desta célula.
cancer.show()

prob_cancer_given_positive = ...

print('Probability of cancer given positive test result: {}'.format(prob_cancer_given_positive))
```

```python
grader.check("q3_2_2")
```

**Questão 3.2.3.** Olhe para o array completo de ícones novamente. Usando isso, como você calcularia a probabilidade condicional de ter câncer dado um resultado positivo no teste?

Execute a próxima célula de código para ver o array completo de ícones.

_Digite sua resposta aqui, substituindo este texto._

```python
# O array completo de ícones é fornecido aqui para sua conveniência.
# Escreva sua resposta na célula anterior.
display_grouped_icon_array(people, "people who've taken a cancer test")
```

**Questão 3.2.4.** Sua resposta para a questão 3.2.2 é maior do que a proporção geral de pessoas na população que têm câncer (dada como 1% na 3.1.1)? Isso faz sentido? Verifique com seus colegas ou um membro da equipe para ver se você está no caminho certo.

_Digite sua resposta aqui, substituindo este texto._

# 4. Uma Prioridade Subjetiva

Estar certo nem sempre é satisfatório. Classificar um paciente Positivo como não tendo câncer ainda parece de alguma forma errado, para um teste tão preciso. Nesta seção, vamos examinar mais a fundo a base do nosso cálculo de probabilidade: a suposição de aleatoriedade.

Nossa suposição era que uma pessoa escolhida aleatoriamente foi testada de nossa população de 10.000 indivíduos. Isso não acontece na realidade. As pessoas vão fazer o teste porque pensam que podem ter a doença, ou porque seu médico acha que elas podem ter a doença. As pessoas que fazem o teste não são membros escolhidos aleatoriamente da população.

É por isso que nossa intuição sobre as pessoas sendo testadas não estava se encaixando bem com a resposta que obtivemos. Na **Questão 3.2.2**, calculamos que a probabilidade de ter câncer dado um resultado positivo no teste era de **0,3125**, ou **31,25%**. Isso parece incomumente baixo, especialmente dado o que sabemos sobre a eficácia dos rastreamentos de câncer. Estávamos imaginando uma situação realista de um paciente indo fazer o teste porque havia algum motivo para isso, enquanto o cálculo foi baseado em uma pessoa escolhida aleatoriamente sendo testada.

Então, vamos refazer nosso cálculo sob a suposição mais realista de que o paciente está fazendo o teste porque o médico acha que há uma chance de o paciente ter a doença e comparar as duas probabilidades.

**Observação:** Para mais informações sobre a Parte 4 deste laboratório, confira a [Seção 18.2.2](https://inferentialthinking.com/chapters/18/2/Making_Decisions.html#a-subjective-prior) do livro didático.

## 4.1. Uma Nova Visão sobre a População
Suponha agora que nossa população de 10.000 indivíduos leva em conta o fato de que escolher fazer um teste médico não é feito ao acaso; deve haver algum motivo para um indivíduo fazer um. Assim, nossa população modificada tem as seguintes propriedades:
1. 1.000 pessoas têm câncer.
2. Entre as 1.000 pessoas que têm câncer, 900 têm resultados positivos em um teste de câncer e 100 têm resultados negativos.
3. As outras 9.000 pessoas não têm câncer.
4. Entre essas 9.000 pessoas, 180 têm resultados positivos em um teste de câncer e as outras 8.820 têm resultados negativos. (Então, 180 veem resultados "falsos positivos".)

Abaixo, geramos uma tabela da população chamada `people_new` e uma tabela de contingência com dados dessas 10.000 pessoas hipotéticas chamada `cancer_new`.

```python
# Apenas execute esta célula para carregar a tabela
people_new = Table().with_columns(
    "status", ["cancer", "cancer", "no cancer", "no cancer"],
    "test status", ["positive", "negative", "positive", "negative"],
    "count", [900, 100, 180, 8820])
people_new
```

```python
# Apenas execute esta célula para carregar a tabela
cancer_new = Table().with_columns(
    "status", ["cancer", "no cancer"],
    "negative", [100, 8820],
    "positive", [900, 180])
cancer_new
```

Como antes, podemos apresentar os dados acima em um array de ícones. Aqui estão os dados de `people_new` nesse formato.

```python
display_grouped_icon_array(people_new, "people who've taken a cancer test")
```

## 4.2. Interpretação e Comparação
Como na Parte 3, vamos supor que a Pessoa X tenha um resultado de teste positivo. Isso significa que agora você pode restringi-la a fazer parte de apenas um dos dois seguintes grupos:
1. As pessoas com câncer que têm um resultado de teste positivo.
2. As pessoas sem câncer que têm um resultado de teste positivo.

Assim como fizemos na Parte 3 deste laboratório, aqui está um array de ícones para esses dois grupos:

```python
# Apenas execute esta célula.
display_grouped_icon_array(people_new.where("test status", are.equal_to("positive")).drop(1), "people who have a positive test result")
```

**Questão 4.2.1.** Usando a nova tabela de contingência, `cancer_new`, que foi fornecida no início desta seção, escreva código para calcular a probabilidade exata de um indivíduo ter câncer dado que ele recebeu um resultado positivo no teste.

Execute a próxima célula de código para ver a tabela `cancer_new`, depois preencha `prob_cancer_given_positive_new` com seu código.

```python
# Execute esta célula primeiro para ver a tabela cancer_new. Depois, preencha a próxima linha desta célula.
cancer_new.show()

prob_cancer_given_positive_new = ...

print(f'Probability of cancer given positive test result: {prob_cancer_given_positive_new}')
```

```python
grader.check("q4_2_1")
```

**Questão 4.2.2.** Como a sua probabilidade na questão 3.2.2 se compara com a sua resposta da 4.2.1? Isso faz sentido? Verifique com seus colegas ou um membro da equipe para ver se você está no caminho certo.

_Digite sua resposta aqui, substituindo este texto._

## 5. Submissão

<img src="harley.jpeg" alt="desenho" width="300"/>

Voltamos ao início! Harley te deu as boas-vindas ao Data 8, e ele está aqui novamente para te despedir. Au!

Parabéns, você terminou o laboratório final!

**Passos importantes para a submissão:** 
1. Execute os testes e verifique se todos passaram.
2. Escolha **Salvar Notebook** no menu **Arquivo**, depois **execute a célula final**. 
3. Clique no link para baixar o arquivo zip.
4. Então, submeta o arquivo zip para a tarefa correspondente de acordo com as instruções do seu instrutor. 

**É sua responsabilidade garantir que seu trabalho esteja salvo antes de executar a última célula.**

## Submissão

Certifique-se de ter executado todas as células no seu caderno em ordem antes de executar a célula abaixo, para que todas as imagens/gráficos apareçam na saída. A célula abaixo irá gerar um arquivo zip para você submeter. **Por favor, salve antes de exportar!**

```python
# Salve seu caderno primeiro, depois execute esta célula para exportar sua submissão.
grader.export(pdf=False, run_tests=True)
```

Tests for q2_0_1 (points: None)
```
>>> 0 < probability_large_shiny < 1 
True
```
```
>>> 
>>> import hashlib 
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.round(probability_large_shiny, 3))
'8404599d79837400f000c64a4fa1cc0e'
```


Tests for q2_0_2 (points: None)
```
>>> import hashlib 
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.round(probability_shiny, 3))
'a83a7d5356406fb9bdbb4d93697cee38'
```


Tests for q2_1_1 (points: None)
```
>>> 
>>> import hashlib 
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.round(probability_shiny_given_large, 3))
'57eeec0a6974ecb4e9fcf68fab052f7b'
```


Tests for q2_1_2 (points: None)
```
>>> 
>>> import hashlib 
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.round(probability_large_given_shiny, 3))
'54fbf38cf649866815e0fefc46a1f6c7'
```


Tests for q3_1_1 (points: None)
```
>>> sum(cancer.column(1)) == 9712
True
```
```
>>> sum(cancer.column(2)) == 288
True
```
```
>>> cancer
status    | negative | positive
cancer    | 10       | 90
no cancer | 9702     | 198
```


Tests for q3_2_1 (points: None)
```
>>> 
>>> import hashlib 
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(rough_prob_cancer_given_positive))
'a87ff679a2f3e71d9181a67b7542122c'
```


Tests for q3_2_2 (points: None)
```
>>> 
>>> import hashlib 
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.round(prob_cancer_given_positive, 4))
'15432ad172f7f319fe502e6337b3f549'
```


Tests for q4_2_1 (points: None)
```
>>> 
>>> import hashlib 
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.round(prob_cancer_given_positive_new, 4))
'4f59b3fd6f5001542919c239391a3278'
```
