```python
# Inicialize o Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

# Laboratório 7: O Grande Concurso Britânico de Panificação (Teste A/B)

Bem-vindo ao Laboratório 7 de Data 8! O laboratório desta semana se concentrará em Testes A/B usando dados do popular programa de televisão britânico, [*The Great British Bakeoff*](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off).

**Submissão**: Quando terminar, execute todas as células, exceto a última, selecione Arquivo > Salvar Notebook e, em seguida, execute a célula final. Em seguida, envie o arquivo zip baixado, que inclui seu notebook, de acordo com as instruções do seu instrutor.

Primeiro, configure o notebook executando a célula abaixo.

```python
# Execute esta célula para configurar o notebook, mas por favor não a altere.

# Estas linhas importam os módulos Numpy e Datascience.
import numpy as np
from datascience import *

# Estas linhas fazem algumas mágicas de plotagem.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', (FutureWarning, np.VisibleDeprecationWarning))

import d8error
```

## 1. Teste A/B

O teste A/B é uma forma de teste de hipóteses que permite fazer comparações entre duas distribuições. Também podemos nos referir a um teste A/B como um teste de permutação.

Quase nunca será solicitado explicitamente que você realize um teste A/B. Certifique-se de que pode identificar situações em que o teste é apropriado e saiba como implementar corretamente cada etapa. Muitas vezes, usamos um teste A/B para determinar se duas amostras vieram da mesma distribuição subjacente.

**Questão 1.1.** As seguintes afirmações são os passos de um teste de hipóteses A/B apresentados em uma *ordem aleatória*:

1. Escolha uma estatística de teste (normalmente a diferença nas médias entre duas categorias)

2. Embaralhe os rótulos da amostra original, encontre sua estatística de teste simulada e repita muitas vezes

3. Encontre o valor da estatística de teste observada

4. Calcule o valor p com base em suas estatísticas de teste observadas e simuladas

5. Defina um modelo nulo e alternativo

6. Use o valor p e o limite de valor p para tirar uma conclusão sobre a hipótese nula

Atribua `ab_test_order` a um array de inteiros que contém a ordem correta de um teste A/B, onde o primeiro item do array é o primeiro passo de um teste A/B e o último item do array é o último passo de um teste A/B.

```python
ab_test_order = ...
```

```python
grader.check("q1_1")
```

**Questão 1.2.** Se a hipótese nula de um teste A/B estiver correta, a ordem dos rótulos deveria afetar as diferenças nas médias entre cada grupo? Por que embaralhamos os rótulos em um teste A/B? Se você estiver em uma seção de laboratório, confirme sua resposta com um colega ou membro da equipe antes de prosseguir.

_Digite sua resposta aqui, substituindo este texto._

## 2. O Grande Concurso Britânico de Panificação

>"O Grande Concurso Britânico de Panificação (frequentemente abreviado para Bake Off ou GBBO) é uma competição britânica de panificação na televisão, produzida pela Love Productions, na qual um grupo de padeiros amadores compete entre si em uma série de rodadas, tentando impressionar um grupo de jurados com suas habilidades de panificação" [Wikipedia](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off)

Para cada semana da competição, os jurados atribuem a um concorrente o título de "Padeiro Estrela". No final, um vencedor é coroado a cada temporada. Usando essas informações, gostaríamos de investigar como ganhar prêmios de Padeiro Estrela afeta as chances de vencer uma temporada do programa.

**Questão 2.1.** Queremos saber se ganhar mais prêmios de Padeiro Estrela *causa* uma mudança na probabilidade de vencer a temporada. Por que não é suficiente comparar as taxas de padeiro estrela para vencedores e perdedores?

_Digite sua resposta aqui, substituindo este texto._

### Executando um Experimento

Vamos realizar o seguinte teste de hipóteses para determinar a associação entre ganhar e o número de prêmios de Padeiro Estrela. A população que estamos examinando é todos os concorrentes das temporadas 2 a 11 do GBBO. Vamos usar as seguintes hipóteses nula e alternativa:

**Hipótese nula:** A distribuição de prêmios de Padeiro Estrela entre os concorrentes que ganharam sua temporada e os concorrentes que não ganharam sua temporada é a mesma.

**Hipótese alternativa:** Os concorrentes que ganham sua temporada do programa ganharão mais prêmios de Padeiro Estrela em média.

Nossa hipótese alternativa está relacionada à nossa suspeita de que os concorrentes que ganham mais prêmios de Padeiro Estrela são mais habilidosos, então eles têm mais chances de ganhar a temporada.

**Questão 2.2.** Devemos usar um teste A/B para testar essas hipóteses? Se sim, qual é o nosso grupo "A" e qual é o nosso grupo "B"?

_Digite sua resposta aqui, substituindo este texto._

Verifique suas respostas com seus colegas ou um membro da equipe antes de passar para a próxima seção.

A tabela `bakers` abaixo descreve o número de prêmios de Padeiro Estrela que cada concorrente ganhou e se eles ganharam ou não a sua temporada (`1` se ganharam, `0` se não ganharam). Os dados foram agregados manualmente da Wikipedia para as temporadas 2-11 do programa. Randomizamos a ordem das linhas para não estragar o resultado do programa.

```python
bakers = Table.read_table("star_bakers.csv")
bakers.show(3)
```

**Questão 2.3.** Crie uma nova tabela chamada `means` que contém a média do número de prêmios de Padeiro Estrela para os padeiros que não ganharam (`won==0`) e para os padeiros que ganharam (`won==1`). A tabela deve ter os nomes das colunas `won` e `star baker awards mean`.

```python
means = ...
means
```

```python
grader.check("q2_3")
```

**Questão 2.4.** Visualize a distribuição dos prêmios de Padeiro Estrela para vencedores e não vencedores. Você deve usar os bins que fornecemos.

Dica: Você vai querer usar o argumento de grupo de `tbl.hist`. Para produzir vários histogramas sobrepostos com base em valores únicos em uma determinada coluna, podemos fazer algo como `tbl.hist(..., group=<col_name>, bins=...)`!

```python
useful_bins = np.arange(0, 7)
...
```

**Questão 2.5.** Queremos descobrir se há uma diferença na distribuição de prêmios de Padeiro Estrela entre vencedores e não vencedores.

Qual deve ser a estatística de teste? Quais valores desta estatística de teste apoiam a nula, e quais valores apoiam a alternativa?

Se você estiver no laboratório, confirme sua resposta com um colega ou membro da equipe antes de prosseguir.

Dica: Você deve pensar sobre quais medidas usamos para descrever uma distribuição.

_Digite sua resposta aqui, substituindo este texto._

**Questão 2.6.** Defina `observed_difference` como a estatística de teste observada usando a tabela `means`.

```python
observed_difference = ...
observed_difference
```

```python
grader.check("q2_6")
```

**Questão 2.7.** Dada uma tabela como `bakers`, uma coluna de rótulos `label_col` e uma coluna de valores `val_col`, escreva uma função que calcule a estatística de teste apropriada.

*Dica:* Certifique-se de levar em conta a direcionalidade de nossa hipótese alternativa.

```python
def find_test_stat(tbl, label_col, val_col):
    ...

find_test_stat(bakers, "won", "star baker awards")
```

```python
grader.check("q2_7")
```

Quando executamos uma simulação para teste A/B, reamostramos **embaralhando os rótulos** da amostra original. Se a hipótese nula for verdadeira e as distribuições de prêmios de Padeiro Estrela forem as mesmas, esperamos que a diferença na média de prêmios de Padeiro Estrela não mude quando os rótulos `"won"` forem alterados.

**Questão 2.8.** Escreva uma função `simulate_and_test_statistic` para calcular um teste do nosso teste A/B. Sua função deve executar uma simulação e retornar uma estatística de teste.

```python
def simulate_and_test_statistic(tbl, labels_col, values_col):
    ...

simulate_and_test_statistic(bakers, "won", "star baker awards")
```

```python
grader.check("q2_8")
```

**Questão 2.9.** Simule 5000 testes do nosso teste A/B e armazene as estatísticas de teste em um array chamado `differences`.

```python
# Esta célula pode levar alguns segundos para executar
differences = make_array()

...
                                                 
differences
```

```python
grader.check("q2_9")
```

Execute a célula abaixo para visualizar um histograma de suas estatísticas de teste simuladas plotadas com sua estatística de teste observada.

```python
Table().with_column('Difference Between Group Means', differences).hist(bins=20)
plots.scatter(observed_difference, 0, color='red', s=30, zorder=2)
plots.ylim(-0.1, 1.35);
```

**Questão 2.10.** Encontre o valor p para o seu teste e atribua-o a `empirical_p`.

```python
empirical_p = ...
empirical_p
```

```python
grader.check("q2_10")
```

**Questão 2.11.** Usando um limite de valor p de 5%, tire uma conclusão sobre as hipóteses nula e alternativa. Descreva suas descobertas usando uma linguagem simples e não técnica. O que sua análise lhe diz sobre a associação entre prêmios de Padeiro Estrela e vitória? O que você pode afirmar sobre a causalidade a partir de sua análise estatística? Chame seu GSI ou AI e confirme sua resposta.

_Digite sua resposta aqui, substituindo este texto._

```python
grader.check("q2_10")
```

## 3. Submissão

Minnie queria dizer parabéns por terminar o Laboratório 7!! Agora você pode ser como a Minnie e aproveitar suas férias de primavera!

<img src="minnie_car.jpg" alt="drawing" width="300"/>

<img src="minnie_beach.jpg" alt="drawing" width="300"/>

**Passos importantes para a submissão:** 
1. Execute os testes e verifique se todos passaram.
2. Escolha **Salvar Notebook** no menu **Arquivo**, então **execute a última célula**. 
3. Clique no link para baixar o arquivo zip.
4. Em seguida, envie o arquivo zip para a tarefa correspondente de acordo com as instruções do seu instrutor. 

**É sua responsabilidade garantir que seu trabalho esteja salvo antes de executar a última célula.**

## Submissão

Certifique-se de ter executado todas as células em seu notebook em ordem antes de executar a célula abaixo, para que todas as imagens/gráficos apareçam na saída. A célula abaixo irá gerar um arquivo zip para você enviar. **Por favor, salve antes de exportar!**

```python
# Primeiro salve seu notebook, depois execute esta célula para exportar sua submissão.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> len(ab_test_order) == 6
True
```
```
>>> np.allclose(np.array(ab_test_order)[:3] % 2 == 1, True) # Check the first three elements in your array.
True
```
```
>>> np.allclose(np.array(ab_test_order)[3:] % 2 == 0, True) # Check the last three elements in your array.
True
```
```
>>> 
>>> # This imports a hashing library for the autograder.
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.array(ab_test_order).astype(int))
'a7196ed0f271c873d9750cb92422d911'
```


Tests for q2_10 (points: None)
```
>>> 0 <= empirical_p < 0.05
True
```


Tests for q2_3 (points: None)
```
>>> means.num_rows
2
```
```
>>> np.round(min(means.column("star baker awards mean")), 2) == 0.65
True
```
```
>>> np.round(max(means.column("star baker awards mean")), 2) == 1.5
True
```


Tests for q2_6 (points: None)
```
>>> isinstance(observed_difference, float)
True
```
```
>>> float(round(observed_difference, 3))
0.848
```


Tests for q2_7 (points: None)
```
>>> np.isclose(round(find_test_stat(bakers, "won", "star baker awards"), 3) - 0.848, 0)
True
```


Tests for q2_8 (points: None)
```
>>> test_stat = round(simulate_and_test_statistic(bakers, "won", "star baker awards"), 3)
>>> -2 < test_stat < 2
True
```
```
>>> np.random.seed(1)
>>> test_stat2 = simulate_and_test_statistic(bakers, "won", "star baker awards")
>>> np.round(test_stat2, 3) == -0.023 or np.round(test_stat2, 3) == -0.132
True
```


Tests for q2_9 (points: None)
```
>>> len(differences)
5000
```
```
>>> abs(np.average(differences)) < 0.05 # On average, your test statistic should be close to 0
True
```
```
>>> all(differences == differences.item(0)) == False # Make sure all of the test statistics are different
True
```
