```python
# Inicialize o otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Tarefa 1: Causalidade e Expressões

Por favor, complete este notebook preenchendo as células fornecidas. Antes de começar, execute a célula anterior para carregar os testes fornecidos.

**Leituras Recomendadas:**

- [O que é Ciência de Dados?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causalidade e Experimentos](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programação em Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

Para todos os problemas nos quais você deve escrever explicações e frases, você **deve** fornecer sua resposta no espaço designado. Além disso, ao longo desta tarefa e de todas as futuras, certifique-se de não reatribuir variáveis ao longo do notebook! Por exemplo, se você usar `max_temperature` em sua resposta para uma pergunta, não o reatribua posteriormente. Caso contrário, você falhará nos testes que pensou ter passado anteriormente!


**Observação: Esta tarefa possui testes ocultos. Isso significa que, mesmo que os testes digam que 100% foram aprovados, isso não significa que sua nota final será 100%. Iremos executar mais testes ocultos para corretude assim que todos entregarem a tarefa.**

Compartilhar respostas diretamente não é permitido, mas discutir problemas com a equipe do curso ou com outros estudantes é incentivado.

Você deve começar cedo para ter tempo de obter ajuda se estiver com dificuldades.

## 1. Aritmética Assustadora

<!-- BEGIN QUESTION -->

Um anúncio para os Sistemas de Segurança ADT diz:

> "Quando você sai de férias, os ladrões vão trabalhar [...] De acordo com estatísticas do FBI, mais de 25% dos arrombamentos domiciliares ocorrem entre o Memorial Day e o Labor Day."

Os dados no anúncio sustentam a afirmação de que os ladrões têm mais probabilidade de trabalhar durante o período entre o Memorial Day e o Labor Day? Por favor, explique sua resposta. **(6 Pontos)**

**Nota:** Você pode assumir que "mais de 25%" significa apenas ligeiramente acima. Se fosse muito mais alto, digamos mais próximo de 30%, então os profissionais de marketing teriam dito isso.

**Nota:** O Memorial Day é observado no último Monday de maio e o Labor Day é observado no primeiro Monday de setembro.



_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

## 2. Personagens em Little Women


Na aula, contamos o número de vezes que os personagens literários foram nomeados em cada capítulo do livro clássico [*Little Women*](https://inferentialthinking.com/chapters/01/3/1/Literary_Characters.html?highlight=little%20women). Na ciência da computação, a palavra "caractere" também se refere a uma letra, dígito, espaço ou sinal de pontuação; qualquer elemento individual de um texto. O código a seguir gera um gráfico de dispersão no qual cada ponto corresponde a um capítulo de *Mulherzinhas*. A posição horizontal de um ponto mede o número de pontos finais no capítulo. A posição vertical mede o número total de caracteres.

```python
# Apenas execute esta célula.

# Esta célula contém código que ainda não foi abordado no curso,
# mas você deverá ser capaz de interpretar o gráfico de dispersão que ele gera.

from datascience import *
from urllib.request import urlopen
import numpy as np
import d8error
%matplotlib inline

little_women_url = 'https://www.inferentialthinking.com/data/little_women.txt'
chapters = urlopen(little_women_url).read().decode().split('CHAPTER ')[1:]
text = Table().with_column('Chapters', chapters)
Table().with_columns(
    'Periods',    np.char.count(chapters, '.'),
    'Characters', text.apply(len, 0)
    ).scatter(0)
```

**Questão 1.** Aproximadamente quantos pontos finais existem no capítulo com mais caracteres? Atribua 1, 2, 3, 4 ou 5 ao nome `characters_q1` abaixo. **(4 Pontos)**

1. 250
2. 390
3. 440
4. 32.000
5. 40.000

**Nota:** Se você encontrar um erro `NameError: name 'grader' is not defined` na célula de autograder abaixo (e em qualquer atribuição), por favor, execute novamente a primeira célula no topo deste notebook!


```python
characters_q1 = ...
```

```python
grader.check("q2_1")
```

O teste acima verifica se suas respostas estão no formato correto. **Este teste não verifica se você respondeu corretamente**, apenas se atribuiu um número com sucesso em cada célula de resposta de múltipla escolha.

**Questão 2.** Qual dos seguintes capítulos tem mais caracteres por ponto final? Atribua 1, 2 ou 3 ao nome `characters_q2` abaixo. **(4 Pontos)**

1. O capítulo com cerca de 60 pontos finais
2. O capítulo com cerca de 350 pontos finais
3. O capítulo com cerca de 440 pontos finais


```python
characters_q2 = ...
```

```python
grader.check("q2_2")
```

Mais uma vez, o teste acima verifica se suas respostas estão no formato correto, mas não se você respondeu corretamente.

Para descobrir mais fatos interessantes deste gráfico, confira [Seção 1.3.2](https://inferentialthinking.com/chapters/01/3/2/Another_Kind_Of_Character.html) no livro didático.

## 3. Nomes e Declarações de Atribuição

**Questão 1.** Quando você executa a célula a seguir, o Python produz uma mensagem de erro críptica.

```python
4 = 2 + 2
```

Escolha a melhor explicação do que há de errado com o código e, em seguida, atribua 1, 2, 3 ou 4 ao `names_q1` abaixo para indicar sua resposta. **(4 Pontos)**

1. Python é inteligente e já sabe que `4 = 2 + 2`.

2. Em Python, é uma regra que o sinal `=` deve ter um nome de variável à sua esquerda, e `4` não é um nome de variável.

3. Deveria ser `2 + 2 = 4`.

4. Não recebo uma mensagem de erro. Esta é uma pergunta de pegadinha.

```python
names_q1 = ...
```

```python
grader.check("q3_1")
```

**Questão 2.** Quando você executa a célula a seguir, o Python produzirá outra mensagem de erro críptica.

```python
two = 3
six = two plus two
```

Escolha a melhor explicação do que há de errado com o código e atribua 1, 2, 3 ou 4 ao `names_q2` abaixo para indicar sua resposta. **(4 Pontos)**

1. A operação `plus` só se aplica a números, não à palavra "two".

2. O nome "two" não pode ser atribuído ao número 3.

3. Dois mais dois são quatro, não seis.

4. O nome `plus` não é um operador embutido; em vez disso, a adição usa `+`.

```python
names_q2 = ...
```

```python
grader.check("q3_2")
```

**Questão 3.** Execute a célula a seguir.

```python
x = 2
y = 3 * x
x = 4
```

Qual é o valor de `y` após a execução desta célula e por quê? Escolha a melhor explicação e atribua 1, 2, 3 ou 4 ao `names_q3` abaixo para indicar sua resposta. **(4 Pontos)**

1. `y` é igual a 6, porque o segundo `x = 4` não tem efeito uma vez que `x` já estava definido.

2. `y` é igual a 6, porque `x` era 2 quando `y` foi atribuído, e 3 * 2 é 6.

3. `y` é igual a 12, porque `x` é 4 e 3 * 4 é 12.

4. `y` é igual a 12, porque atribuir `x` para 4 atualizará `y` para 12, já que `y` foi definido em termos de `x`.

```python
names_q3 = ...
```

```python
grader.check("q3_3")
```

## 4. Diferenças Entre Cursos

O Escritório de Planejamento e Análise da Universidade de Berkeley fornece dados sobre vários aspectos do campus. Adaptada do site da OPA, a tabela abaixo exibe o número de formandos em três cursos nos anos acadêmicos de 2008-2009 e 2017-2018.

| Curso                              | 2008-2009    | 2017-2018   |
|------------------------------------|--------------|-------------|
| Estudos de Gênero e Mulheres       |      17      |    28       |
| Linguística                        |      49      |    67       |
| Retórica                           |      113     |    56       |

**Questão 1.** Suponha que você queira encontrar a maior diferença absoluta entre o número de formandos nos dois anos, entre os três cursos.

Na célula abaixo, calcule esse valor e chame-o de `biggest_change`. Use uma única expressão (uma única linha de código) para calcular a resposta. Deixe o Python realizar toda a aritmética (como subtrair 49 de 67) em vez de simplificar a expressão você mesmo. A função embutida `abs` recebe uma entrada numérica e retorna o valor absoluto. A função embutida `max` pode receber 3 argumentos e retorna o máximo dos três números. **(5 Pontos)**

```python
biggest_change = ...
biggest_change
```

```python
grader.check("q4_1")
```

**Questão 2.** Qual dos três cursos teve a **menor** diferença absoluta? Atribua `smallest_change_major` como 1, 2 ou 3, onde cada número corresponde ao seguinte curso:

1. Estudos de Gênero e Mulheres  
2. Linguística  
3. Retórica

Escolha o número que corresponde ao curso com a menor diferença absoluta.

Você deve ser capaz de responder por meio de um cálculo mental aproximado, sem precisar calcular o valor exato para cada curso. **(4 Pontos)** 

```python
smallest_change_major = ...
smallest_change_major
```

```python
grader.check("q4_2")
```

**Questão 3.** Para cada grande área, defina a "mudança relativa" da seguinte maneira: $\large{\frac{\text{diferença absoluta}}{\text{valor em 2008-2009}} * 100}$

Preencha o código abaixo de modo que `gws_relative_change`, `linguistics_relative_change` e `rhetoric_relative_change` sejam atribuídos às mudanças relativas de suas respectivas áreas principais. **(5 Pontos)**

```python
gws_relative_change = (abs(...) / 17) * 100
linguistics_relative_change = ...
rhetoric_relative_change = ...
gws_relative_change, linguistics_relative_change, rhetoric_relative_change
```

```python
grader.check("q4_3")
```

**Questão 4.** Atribua `biggest_rel_change_major` como 1, 2 ou 3, onde cada número corresponde ao seguinte:

1. Estudos de Gênero e Mulheres
2. Linguística
3. Retórica

Escolha o número que corresponde à área principal com a maior mudança relativa. **(4 Pontos)**

```python
biggest_rel_change_major = ...
biggest_rel_change_major
```

```python
grader.check("q4_4")
```

## 5. Estudo sobre Miopia

[Miopia](https://pt.wikipedia.org/wiki/Miopia) resulta de uma série de fatores genéticos e ambientais. Em 1999, Quinn et al. estudaram a relação entre miopia e iluminação ambiente à noite (por exemplo, de luzes noturnas ou luzes de ambiente) durante a infância.

<!-- BEGIN QUESTION -->

**Questão 1.** Os dados foram coletados pelo seguinte procedimento, relatado no estudo. "Entre janeiro e junho de 1998, os pais de crianças com idades entre 2 e 16 anos [...] que foram atendidas como pacientes ambulatoriais em uma clínica de oftalmologia pediátrica universitária completaram um questionário sobre a exposição à luz da criança tanto no presente quanto antes dos 2 anos de idade." Este estudo foi observacional ou foi um experimento controlado? Explique. **(5 Pontos)**


_Escreva sua resposta aqui. Substituindo este texto._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Questão 2.** O estudo descobriu que, das crianças que dormiram com a luz do quarto acesa antes dos 2 anos, 55% eram míopes. Das crianças que dormiram com uma luz noturna antes dos 2 anos, 34% eram míopes. Das crianças que dormiram no escuro antes dos 2 anos, 10% eram míopes. O estudo concluiu o seguinte: "A prevalência de miopia [...] durante a infância estava fortemente associada à exposição à luz ambiente durante o sono à noite nos primeiros dois anos após o nascimento."

Os dados suportam esta afirmação? Por quê ou por que não? Você pode interpretar "fortemente" de qualquer maneira qualitativa razoável. **(5 Pontos)**


_Escreva sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Questão 3.** Em 13 de maio de 1999, a CNN relatou os resultados deste estudo sob a manchete, "Luz noturna pode levar à miopia." A conclusão do estudo afirma que a luz noturna causa miopia? **(5 Pontos)**


_Escreva sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Questão 4.** O último parágrafo do relatório da CNN mencionou que "vários especialistas em olhos" apontaram que o estudo deveria ter levado em conta a hereditariedade.

A miopia é transmitida de pais para filhos. Pais míopes têm maior probabilidade de ter filhos míopes e podem também ter maior probabilidade de deixar as luzes acesas habitualmente (já que os pais têm visão ruim). De que maneira o conhecimento dessa possível ligação genética afeta a forma como interpretamos os dados do estudo? Explique. **(5 Pontos)**


_Escreva sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

## 6. Estudando os Sobreviventes

O Reverendo Henry Whitehead estava cético em relação à conclusão de John Snow sobre a bomba de água da Broad Street. Após o término da epidemia de cólera na Broad Street, Whitehead tentou provar que Snow estava errado. (A história do evento está detalhada [aqui](http://www.ncbi.nlm.nih.gov/pmc/articles/PMC1034367/pdf/medhist00183-0026.pdf).)

Ele percebeu que Snow havia focado sua análise quase inteiramente naqueles que haviam morrido. Whitehead, portanto, investigou os hábitos de consumo de água das pessoas na área da Broad Street que não morreram durante o surto.

Qual é a principal razão pela qual era importante estudar esse grupo? Atribua 1, 2 ou 3 ao nome `survivor_answer` abaixo. **(4 Pontos)**

1. Se Whitehead tivesse descoberto que muitas pessoas beberam água da bomba da Broad Street e não contraíram cólera, isso seria uma evidência contra a hipótese de Snow.

2. Os sobreviventes poderiam fornecer informações adicionais sobre o que mais poderia ter causado a cólera, potencialmente descobrindo outra causa.

3. Ao considerar os sobreviventes, Whitehead poderia ter identificado uma cura para a cólera.

```python
survivor_answer = ...
```

```python
grader.check("q6_1")
```

**Nota:** Whitehead acabou encontrando mais provas de que a bomba da Broad Street desempenhou um papel central na disseminação da doença entre as pessoas que viviam nas proximidades. Eventualmente, ele se tornou um dos maiores defensores de Snow.

Você concluiu a Tarefa de Casa 1!

**Passos importantes para a submissão:**
1. Execute os testes e verifique se todos estão passando.
2. Escolha **Salvar Notebook** no menu **Arquivo**, em seguida, **execute a célula final**.
3. Clique no link para baixar o arquivo zip.
4. Em seguida, envie o arquivo zip para a tarefa correspondente de acordo com as instruções do seu instrutor.

**É sua responsabilidade garantir que seu trabalho esteja salvo antes de executar a última célula.**

## Submissão

Certifique-se de ter executado todas as células em seu notebook na ordem correta antes de rodar a célula abaixo, para que todas as imagens/gráficos apareçam na saída. A célula abaixo irá gerar um arquivo zip para você enviar. **Por favor, salve antes de exportar!**

```python
# Primeiro, salve seu notebook e depois execute esta célula para exportar sua submissão.
grader.export(pdf=False, run_tests=True)
```

Tests for q2_1 (points: [0])
```
>>> 1 <= characters_q1 <= 5
True
```


Tests for q2_2 (points: [0])
```
>>> 1 <= characters_q2 <= 3
True
```


Tests for q3_1 (points: [0])
```
>>> 1 <= names_q1 <= 4
True
```


Tests for q3_2 (points: [0])
```
>>> 1 <= names_q2 <= 4
True
```


Tests for q3_3 (points: [0])
```
>>> 1 <= names_q3 <= 4
True
```


Tests for q4_1 (points: [0])
```
>>> isinstance(biggest_change, (int, float))
True
```


Tests for q4_2 (points: [0])
```
>>> isinstance(smallest_change_major, (int, float))
True
```


Tests for q4_3 (points: [0, 0, 0])
```
>>> isinstance(gws_relative_change, (int, float))
True
```
```
>>> isinstance(linguistics_relative_change, (int, float))
True
```
```
>>> isinstance(rhetoric_relative_change, (int, float))
True
```


Tests for q4_4 (points: [0])
```
>>> isinstance(biggest_rel_change_major, (int, float))
True
```


Tests for q6_1 (points: [0])
```
>>> 1 <= survivor_answer <= 3
True
```
