```python
# Inicialize o Otter
import otter
grader = otter.Notebook("hw09.ipynb")
```

# Tarefa de Casa 9: Tamanhos de Amostra e Intervalos de Confiança

**Recurso Útil:**

- [Referência Python](http://data8.org/sp22/python-reference.html): Folha de dicas de métodos úteis de arrays e tabelas usados no Data 8!

**Leituras Recomendadas**: 

* [Estimação](https://inferentialthinking.com/chapters/13/Estimation.html)
* [Por Que a Média Importa](https://inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html)

Por favor, complete este caderno preenchendo as células fornecidas. Antes de começar, execute a célula a seguir para configurar o caderno importando algumas bibliotecas úteis. Cada vez que você iniciar seu servidor, precisará executar esta célula novamente.

Para todos os problemas nos quais você deve escrever explicações e frases, você **deve** fornecer sua resposta no espaço designado. **Além disso, ao longo desta tarefa de casa e de todas as futuras, certifique-se de não reatribuir variáveis ao longo do caderno!** Por exemplo, se você usar `max_temperature` na sua resposta para uma questão, não a reatribua mais tarde. Caso contrário, você falhará em testes que pensava ter passado anteriormente!

**Nota: Esta tarefa de casa possui testes ocultos. Isso significa que, mesmo que os testes digam que 100% passaram, não significa que sua nota final será 100%. Nós rodaremos mais testes para correção uma vez que todos entreguem a tarefa de casa.**

Compartilhar diretamente as respostas não é permitido, mas discutir os problemas com a equipe do curso ou com outros estudantes é encorajado.

Você deve começar cedo para que tenha tempo de obter ajuda se estiver com dificuldades.

```python
# Não altere esta célula; apenas execute-a.
import numpy as np
from datascience import *

# Essas linhas fazem alguma mágica sofisticada de plotagem.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## 1. Limitando a Cauda de uma Distribuição

Uma comunidade tem uma idade média de 45 anos com um desvio padrão de 5 anos.

Em cada parte abaixo, preencha o espaço em branco com um percentual que torne a afirmação verdadeira **sem mais suposições**, e explique sua resposta.

*Nota:* Não será concedido crédito para limites vagos como "pelo menos 0%" ou "no máximo 100%". Dê a melhor resposta possível com as informações fornecidas.

<!-- BEGIN QUESTION -->

**Questão 1.1.** Pelo menos _______% das pessoas têm entre 25 e 65 anos de idade. **(6 Pontos)**

_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.2.** At most _______% of the people have ages that are not in the range 25 years to 65 years. **(6 Points)**


_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Questão 1.3.** No máximo _______% das pessoas têm mais de 65 anos de idade. **(6 Pontos)**

*Dica:* Se você estiver com dificuldades, tente imaginar como a distribuição pode parecer neste caso.

_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

## 2. Tamanho da Amostra e Nível de Confiança

Uma turma de ciência de dados na grande Universidade Data 8 quer estimar o percentual de usuários do Facebook entre os estudantes da escola. Para fazer isso, eles precisam coletar uma amostra aleatória de estudantes. Você pode assumir que o método de amostragem deles é equivalente a sortear aleatoriamente com reposição entre os estudantes da escola.

***Por favor, revise a [Seção 14.6](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html#) do livro didático antes de prosseguir com esta seção. Seu trabalho será muito mais rápido dessa forma!***

**Questão 2.1.** Atribua a `smallest` o menor número de estudantes que eles devem amostrar para garantir que um intervalo de confiança de **95%** para o parâmetro tenha uma largura de no máximo 6% de uma extremidade à outra. **(6 pontos)**

*Dica:* Como nossos dados podem ser representados para mostrar se um estudante na amostra é um usuário do Facebook? Dado isso, quais suposições podemos fazer sobre o DP da população? A [Seção 14.6](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html#) pode ser útil!

*Nota:* Embora o verdadeiro menor tamanho de amostra tenha que ser um inteiro, por favor, deixe sua resposta em formato decimal para os nossos testes.

```python
smallest = ...
smallest
```

```python
grader.check("q2_1")
```

<!-- BEGIN QUESTION -->

**Questão 2.2.** Suponha que a turma de ciência de dados decida construir um intervalo de confiança de 90% em vez de um intervalo de 95%, mas ainda exija que a largura do intervalo não seja maior que 6% da extremidade esquerda à direita. Eles precisarão do mesmo tamanho de amostra que em 2.1? Escolha a resposta correta e explique mais sem cálculo. **(6 Pontos)**

1. Sim, eles devem usar o mesmo tamanho de amostra.
2. Não, um tamanho de amostra menor funcionará.
3. Não, eles precisarão de uma amostra maior.

_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

**Questão 2.3.** O professor informa à turma que um intervalo de confiança de 90% para o parâmetro é construído exatamente como um intervalo de 95%, exceto que você precisa ir apenas 1,65 desvios padrão de cada lado da estimativa (±1,65) em vez de 2 desvios padrão de cada lado (±2). Atribua a `smallest_num` o menor número de estudantes que eles devem amostrar para garantir que um intervalo de confiança de **90%** para o parâmetro tenha uma largura de no máximo 6% da extremidade esquerda à direita. **(6 pontos)**

*Nota:* Embora o verdadeiro menor tamanho de amostra tenha que ser um inteiro, por favor, deixe sua resposta em formato decimal para os nossos testes.

```python
smallest_num = ...
smallest_num
```

```python
grader.check("q2_3")
```

Para este próximo exercício, consulte a [Seção 14.3.4](https://inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html#the-standard-normal-cdf) do livro didático para exemplos similares.

Os estudantes estão curiosos sobre como a professora chegou ao valor 1.65 na Questão 2.3. Ela diz que executou as duas células de código a seguir. A primeira chama a função `plot_normal_cdf` da biblioteca `datascience`, que exibe a proporção que é no máximo o número especificado de DP acima da média sob a curva normal plotada com unidades padrão no eixo horizontal. Você pode encontrar a documentação [aqui](http://data8.org/datascience/util.html#datascience.util.plot_normal_cdf).

*Nota:* A sigla `cdf` significa `função de distribuição cumulativa`. Ela mede a proporção à esquerda de um ponto especificado sob um histograma de probabilidade.

```python
plot_normal_cdf(1.65)
```

Para executar a segunda célula, a professora teve que primeiro importar uma biblioteca Python para probabilidade e estatística:

```python
# Apenas execute esta célula.
from scipy import stats
```

Em seguida, ela usou o método `norm.cdf` na biblioteca para encontrar a proporção de ouro acima.

```python
# Apenas execute esta célula.
stats.norm.cdf(1.65)
```

*Nota*: Você não precisa entender como a biblioteca `scipy` funciona ou como usar o método por conta própria.

<!-- BEGIN QUESTION -->

**Questão 2.4.** Isso mostra que a porcentagem em uma distribuição normal que está no máximo 1,65 desvios padrão acima da média é de cerca de **95%**. Explique por que 1,65 é o número correto de desvios padrão a ser usado ao construir um intervalo de confiança de **90%**. **(6 Pontos)**

_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->



```python
# Apenas execute esta célula, não altere.
stats.norm.cdf(2.33)
```

**Questão 2.5.** A célula acima mostra que a proporção que está no máximo 2,33 desvios padrão acima da média em uma distribuição normal é de 99%. Atribua `option` à opção correta para preencher o espaço em branco: **(6 pontos)**

Se você começar na estimativa e for 2,33 desvios padrão para cada lado, então você obterá um intervalo de confiança de _______% para o parâmetro.

1. 99,5
2. 99
3. 98,5
4. 98

```python
option = ...
option
```

```python
grader.check("q2_5")
```

## 3. Polling e a Distribuição Normal

```python
# Não altere esta célula; apenas execute-a.
import numpy as np
from datascience import *

# Essas linhas fazem alguma mágica sofisticada de plotagem.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

Michelle é uma consultora estatística e trabalha para um grupo que apoia a Proposição 68 (que exigiria a rotulagem de todos os eixos horizontais e verticais) chamado Sim na 68. Eles querem saber quantos californianos votarão a favor da proposição.

Michelle faz uma pesquisa com uma amostra aleatória uniforme de todos os eleitores da Califórnia e descobre que 210 dos 400 eleitores amostrados votarão a favor da proposição. Fornecemos uma tabela para você abaixo, que tem 3 colunas: as duas primeiras colunas são idênticas à `sample`. A terceira coluna contém a proporção de eleitores totais que escolheram cada opção.

```python
sample = Table().with_columns(
    "Vote",  make_array("Yes", "No"),
    "Count", make_array(210,   190))

sample_size = sum(sample.column("Count"))
sample_with_proportions = sample.with_column("Proportion", sample.column("Count") / sample_size)
sample_with_proportions
```

**Questão 3.1.** Michelle deseja usar 10.000 reamostragens bootstrap para calcular um intervalo de confiança para a proporção de todos os eleitores da Califórnia que votarão Sim.

Preencha a próxima célula para simular uma distribuição empírica de proporções de Sim. Use reamostragem bootstrap para simular 10.000 resultados de eleição, e atribua `resample_yes_proportions` para conter a proporção de Sim de cada reamostragem bootstrap. Em seguida, visualize `resample_yes_proportions` com um histograma. **Você deve ver uma curva em forma de sino centrada perto da proporção de Sim na amostra original.** **(6 pontos)**

*Dica:* `sample_proportions` pode ser útil aqui!

```python
resample_yes_proportions = make_array()
for i in np.arange(10000):
    resample = ...
    resample_yes_proportions = ...
Table().with_column("Resample Yes proportion", resample_yes_proportions).hist(bins=np.arange(.2, .8, .01))
```

```python
grader.check("q3_1")
```

<!-- BEGIN QUESTION -->

**Questão 3.2.** Por que o Teorema do Limite Central (TLC) se aplica nesta situação, e como ele explica a distribuição que vemos acima? **(6 pontos)**

_Digite sua resposta aqui, substituindo este texto._

<!-- END QUESTION -->

Em uma população cujos membros são 0 e 1, existe uma fórmula simples para o **desvio padrão dessa população**:

$$\text{desvio padrão da população} = \sqrt{(\text{proporção de 0s}) \times (\text{proporção de 1s})}$$

(Descobrir essa fórmula, partindo da definição de desvio padrão, é um exercício divertido para aqueles que gostam de álgebra.)

**Questão 3.3.** Usando apenas o Teorema do Limite Central e os números de eleitores Sim e Não em nossa amostra de 400, *algebricamente* calcule o desvio padrão previsto do array `resample_yes_proportions`. Atribua este número a `approximate_sd`. **Não acesse os dados em `resample_yes_proportions` de forma alguma.** **(6 pontos)**

Lembre-se de que o desvio padrão das médias da amostra pode ser calculado a partir do desvio padrão da população e do tamanho da amostra (a fórmula acima pode ser útil). Se não soubermos o desvio padrão

```python
approx_pop_sd = ...
approximate_sd = ...
approximate_sd
```

```python
grader.check("q3_3")
```

**Questão 3.4.** Calcule o desvio padrão do array `resample_yes_proportions`, que servirá como uma aproximação para o verdadeiro desvio padrão das proporções amostrais possíveis. Isso ajudará a verificar se sua resposta para a questão 3.3 está aproximadamente correta. **(6 pontos)**

```python
exact_sd = ...
exact_sd
```

```python
grader.check("q3_4")
```

**Questão 3.5.** **Novamente, sem acessar `resample_yes_proportions` de forma alguma**, calcule um intervalo de confiança aproximado de 95% para a proporção de eleitores Sim na Califórnia. **(6 pontos)**

A célula abaixo desenha seu intervalo como uma barra vermelha abaixo do histograma de `resample_yes_proportions`; use isso para verificar se sua resposta parece correta.

*Dica:* Quantos desvios padrão correspondem a 95% da distribuição prometida pelo TLC? Relembre a discussão no livro didático [aqui](https://inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html).

*Dica:* A variável `approximate_sd` que você definiu anteriormente pode ser útil!

```python
lower_limit = ...
upper_limit = ...
print('lower:', lower_limit, 'upper:', upper_limit)
```

```python
grader.check("q3_5")
```

```python
# Execute esta célula para plotar seu intervalo de confiança.
Table().with_column("Resample Yes proportion", resample_yes_proportions).hist(bins=np.arange(.2, .8, .01))
plt.plot(make_array(lower_limit, upper_limit), make_array(0, 0), c='r', lw=10);
```

Seu intervalo de confiança deve sobrepor o número 0,5. Isso significa que não podemos ter certeza se a Proposição 68 está ganhando, mesmo que a proporção de Sim na amostra esteja um pouco acima de 0,5.

A campanha Sim na 68 realmente precisa saber se está ganhando. É impossível ter certeza absoluta sem pesquisar toda a população, mas eles ficariam satisfeitos se o desvio padrão da média da amostra fosse apenas 0,005. Eles pedem a Michelle para realizar uma nova pesquisa com um tamanho de amostra grande o suficiente para alcançar isso. (As pesquisas são caras, então a amostra também não deve ser maior do que o necessário.)

Michelle consulta o Capítulo 14 do seu livro didático. Em vez de fazer a suposição conservadora de que o desvio padrão da população é 0,5 (codificando os eleitores Sim como 1 e os eleitores Não como 0), ela decide assumir que é igual ao desvio padrão da amostra,

$$\sqrt{(\text{Proporção de Sim na amostra}) \times (\text{Proporção de Não na amostra})}.$$

Com essa suposição, Michelle decide que uma amostra de 9.975 seria suficiente.

**Questão 3.6.** O tamanho da amostra de Michelle alcança o desvio padrão desejado das médias da amostra? Que desvio padrão você alcançaria com um tamanho de amostra menor? Um tamanho de amostra maior?

Para explorar isso, primeiro calcule o desvio padrão das médias da amostra obtido usando o tamanho da amostra de Michelle. **(6 pontos)**

```python
estimated_population_sd = ...
michelle_sample_size = ...
michelle_sample_mean_sd = ...
print("With Michelle's sample size, you would predict a sample mean SD of %f." % michelle_sample_mean_sd)
```

```python
grader.check("q3_6")
```

**Questão 3.7.** Em seguida, calcule o desvio padrão das médias da amostra que você obteria de um tamanho de amostra menor. Idealmente, você deve escolher um número significativamente menor, mas qualquer tamanho de amostra menor que o de Michelle servirá. **(5 pontos)**

```python
smaller_sample_size = ...
smaller_sample_mean_sd = ...
print("With this smaller sample size, you would predict a sample mean SD of %f" % smaller_sample_mean_sd)
```

```python
grader.check("q3_7")
```

**Questão 3.8.** Finalmente, calcule o desvio padrão das médias da amostra que você obteria de um tamanho de amostra maior. Aqui, um número significativamente maior faria qualquer diferença mais óbvia, mas qualquer tamanho de amostra maior que o de Michelle servirá. **(5 pontos)**

```python
larger_sample_size = ...
larger_sample_mean_sd = ...
print("With this larger sample size, you would predict a sample mean SD of %f" % larger_sample_mean_sd)
```

```python
grader.check("q3_8")
```

**Questão 3.9.** Com base nisso, o tamanho da amostra de Michelle era aproximadamente o mínimo suficiente, dado sua suposição de que o desvio padrão da amostra é o mesmo que o desvio padrão da população? Atribua `min_sufficient` a `True` se 9.975 foi de fato aproximadamente o mínimo suficiente, e `False` se não foi. **(4 pontos)**

```python
min_sufficient = ...
min_sufficient
```

```python
grader.check("q3_9")
```

Você concluiu a Lição de Casa 9!

**Passos importantes para a submissão:**
1. Execute os testes e verifique se todos passaram.
2. Escolha **Salvar Caderno** no menu **Arquivo**, depois **execute a última célula**.
3. Clique no link para baixar o arquivo zip.
4. Então, submeta o arquivo zip para a tarefa correspondente de acordo com as instruções do seu instrutor.

**É sua responsabilidade garantir que seu trabalho esteja salvo antes de executar a última célula.**

## Submissão

Certifique-se de ter executado todas as células no seu caderno em ordem antes de executar a célula abaixo, para que todas as imagens/gráficos apareçam na saída. A célula abaixo irá gerar um arquivo zip para você submeter. **Por favor, salve antes de exportar!**

```python
# Salve seu caderno primeiro, depois execute esta célula para exportar sua submissão.
grader.export(pdf=False, run_tests=True)
```

Tests for q2_1 (points: [0])
```
>>> 1100 <= smallest <= 1150
True
```


Tests for q2_3 (points: [0])
```
>>> 700 <= smallest_num <= 800
True
```


Tests for q2_5 (points: [0])
```
>>> type(option) == int
True
```


Tests for q3_1 (points: [1, 5])
```
>>> len(resample_yes_proportions) == 10000
True
```
```
>>> abs(np.mean(resample_yes_proportions) - .525) < .025 and np.std(resample_yes_proportions) < .08
True
```


Tests for q3_3 (points: [0])
```
>>> approximate_sd < 0.025
True
```


Tests for q3_4 (points: [0])
```
>>> exact_sd < 0.05
True
```


Tests for q3_5 (points: [0])
```
>>> 0.4 <= lower_limit < upper_limit <= 0.7
True
```


Tests for q3_6 (points: [0])
```
>>> type(michelle_sample_size) == int
True
```


Tests for q3_7 (points: [1])
```
>>> smaller_sample_size < michelle_sample_size
True
```


Tests for q3_8 (points: [1])
```
>>> larger_sample_size > michelle_sample_size
True
```


Tests for q3_9 (points: [0])
```
>>> min_sufficient or not min_sufficient 
True
```
